An Advanced Bacterial Single-cell RNA-seq Reveals Biofilm Heterogeneity

Author response:

The following is the authors’ response to the original reviews.

eLife Assessment

The work introduces a valuable new method for depleting the ribosomal RNA from bacterial single-cell RNA sequencing libraries and shows that this method is applicable to studying the heterogeneity in microbial biofilms. The evidence for a small subpopulation of cells at the bottom of the biofilm which upregulates PdeI expression is solid. However, more investigation into the unresolved functional relationship between PdeI and c-di-GMP levels with the help of other genes co-expressed in the same cluster would have made the conclusions more significant.

Many thanks for eLife’s assessment of our manuscript and the constructive feedback. We are encouraged by the recognition of our bacterial single-cell RNA-seq methodology as valuable and its efficacy in studying bacterial population heterogeneity. We appreciate the suggestion for additional investigation into the functional relationship between PdeI and c-di-GMP levels. We concur that such an exploration could substantially enhance the impact of our conclusions. To address this, we have implemented the following revisions: We have expanded our data analysis to identify and characterize genes co-expressed with PdeI within the same cellular cluster (Fig. 3F, G, Response Fig. 10); We conducted additional experiments to validate the functional relationships between PdeI and c-di-GMP, followed by detailed phenotypic analyses (Response Fig. 9B). Our analysis reveals that while other marker genes in this cluster are co-expressed, they do not significantly impact biofilm formation or directly relate to c-di-GMP or PdeI. We believe these revisions have substantially enhanced the comprehensiveness and context of our manuscript, thereby reinforcing the significance of our discoveries related to microbial biofilms. The expanded investigation provides a more thorough understanding of the PdeI-associated subpopulation and its role in biofilm formation, addressing the concerns raised in the initial assessment.

Public Reviews:

Reviewer #1 (Public Review):

Summary:

In this manuscript, Yan and colleagues introduce a modification to the previously published PETRI-seq bacterial single-cell protocol to include a ribosomal depletion step based on a DNA probe set that selectively hybridizes with ribosome-derived (rRNA) cDNA fragments. They show that their modification of the PETRI-seq protocol increases the fraction of informative non-rRNA reads from ~4-10% to 54-92%. The authors apply their protocol to investigating heterogeneity in a biofilm model of E. coli, and convincingly show how their technology can detect minority subpopulations within a complex community.

Strengths:

The method the authors propose is a straightforward and inexpensive modification of an established split-pool single-cell RNA-seq protocol that greatly increases its utility, and should be of interest to a wide community working in the field of bacterial single-cell RNA-seq.

Weaknesses:

The manuscript is written in a very compressed style and many technical details of the evaluations conducted are unclear and processed data has not been made available for evaluation, limiting the ability of the reader to independently judge the merits of the method.

Thank you for your thoughtful and constructive review of our manuscript. We appreciate your recognition of the strengths of our work and the potential impact of our modified PETRI-seq protocol on the field of bacterial single-cell RNA-seq. We are grateful for the opportunity to address your concerns and improve the clarity and accessibility of our manuscript.

We acknowledge your feedback regarding the compressed writing style and lack of technical details, which are constrained by the requirements of the Short Report format in eLife. We have addressed these issues in our revised manuscript as follows:

(1) Expanded methodology section: We have provided a more comprehensive description of our experimental procedures, including detailed protocols for the ribosomal depletion step (lines 435-453) and data analysis pipeline (lines 471-528). This will enable readers to better understand and potentially replicate our methods.

(2) Clarification of technical evaluations: We have elaborated on the specifics of our evaluations, including the criteria used for assessing the efficiency of ribosomal depletion (lines 99-120), and the methods employed for identifying and characterizing subpopulations (lines 155-159, 161-163 and 163-167).

(3) Data availability: We apologize for the oversight in not making our processed data readily available. We have deposited all relevant datasets, including raw and source data, in appropriate public repositories (GEO: GSE260458) and provide clear instructions for accessing this data in the revised manuscript.

(4) Supplementary information: To maintain the concise nature of the main text while providing necessary details, we have included additional supplementary information. This will cover extended methodology (lines 311-318, 321-323, 327-340, 450-453, 533, and 578-589), detailed statistical analyses (lines 492-493, 499-501 and 509-528), and comprehensive data tables to support our findings.

We believe these changes significantly improved the clarity and reproducibility of our work, allowing readers to better evaluate the merits of our method.

Reviewer #2 (Public Review):

Summary:

This work introduces a new method of depleting the ribosomal reads from the single-cell RNA sequencing library prepared with one of the prokaryotic scRNA-seq techniques, PETRI-seq. The advance is very useful since it allows broader access to the technology by lowering the cost of sequencing. It also allows more transcript recovery with fewer sequencing reads. The authors demonstrate the utility and performance of the method for three different model species and find a subpopulation of cells in the E.coli biofilm that express a protein, PdeI, which causes elevated c-di-GMP levels. These cells were shown to be in a state that promotes persister formation in response to ampicillin treatment.

Strengths:

The introduced rRNA depletion method is highly efficient, with the depletion for E.coli resulting in over 90% of reads containing mRNA. The method is ready to use with existing PETRI-seq libraries which is a large advantage, given that no other rRNA depletion methods were published for split-pool bacterial scRNA-seq methods. Therefore, the value of the method for the field is high. There is also evidence that a small number of cells at the bottom of a static biofilm express PdeI which is causing the elevated c-di-GMP levels that are associated with persister formation. Given that PdeI is a phosphodiesterase, which is supposed to promote hydrolysis of c-di-GMP, this finding is unexpected.

Weaknesses:

With the descriptions and writing of the manuscript, it is hard to place the findings about the PdeI into existing context (i.e. it is well known that c-di-GMP is involved in biofilm development and is heterogeneously distributed in several species' biofilms; it is also known that E.coli diesterases regulate this second messenger, i.e. https://journals.asm.org/doi/full/10.1128/jb.00604-15).

There is also no explanation for the apparently contradictory upregulation of c-di-GMP in cells expressing higher PdeI levels. Perhaps the examination of the rest of the genes in cluster 2 of the biofilm sample could be useful to explain the observed association.

Thank you for your thoughtful and constructive review of our manuscript. We are pleased that the reviewer recognizes the value and efficiency of our rRNA depletion method for PETRI-seq, as well as its potential impact on the field. We would like to address the points raised by the reviewer and provide additional context and clarification regarding the function of PdeI in c-di-GMP regulation.

We acknowledge that c-di-GMP’s role in biofilm development and its heterogeneous distribution in bacterial biofilms are well studied. We appreciate the reviewer's observation regarding the seemingly contradictory relationship between increased PdeI expression and elevated c-di-GMP levels. This is indeed an intriguing finding that warrants further explanation.

PdeI is predicted to function as a phosphodiesterase involved in c-di-GMP degradation, based on sequence analysis demonstrating the presence of an intact EAL domain, which is known for this function. However, it is important to note that PdeI also harbors a divergent GGDEF domain, typically associated with c-di-GMP synthesis. This dual-domain structure indicates that PdeI may play complex regulatory roles. Previous studies have shown that knocking out the major phosphodiesterase PdeH in E. coli results in the accumulation of c-di-GMP. Moreover, introducing a point mutation (G412S) in PdeI's divergent GGDEF domain within this PdeH knockout background led to decreased c-di-GMP levels2. This finding implies that the wild-type GGDEF domain in PdeI contributes to maintaining or increasing cellular c-di-GMP levels.

Importantly, our single-cell experiments demonstrated a positive correlation between PdeI expression levels and c-di-GMP levels (Figure 4D). In this revision, we also constructed a PdeI(G412S)-BFP mutation strain. Notably, our observations of this strain revealed that c-di-GMP levels remained constant despite an increase in BFP fluorescence, which serves as a proxy for PdeI(G412S) expression levels (Figure 4D). This experimental evidence, coupled with domain analyses, suggests that PdeI may also contribute to c-di-GMP synthesis, rebutting the notion that it acts solely as a phosphodiesterase. HPLC LC-MS/MS analysis further confirmed that the overexpression of PdeI, induced by arabinose, resulted in increased c-di-GMP levels (Fig. 4E) . These findings strongly suggest that PdeI plays a pivotal role in upregulating c-di-GMP levels.

Our further analysis indicated that PdeI contains a CHASE (cyclases/histidine kinase-associated sensory) domain. Combined with our experimental results showing that PdeI is a membrane-associated protein, we hypothesize that PdeI acts as a sensor, integrating environmental signals with c-di-GMP production under complex regulatory mechanisms.

We understand your interest in the other genes present in cluster 2 of the biofilm and their potential relationship to PdeI and c-di-GMP. Upon careful analysis, we have determined that the other marker genes in this cluster do not significantly impact biofilm formation, nor have we identified any direct relationship between these genes, c-di-GMP, or PdeI. Our focus on PdeI within this cluster is justified by its unique and significant role in c-di-GMP regulation and biofilm formation, as demonstrated by our experimental results. While other genes in this cluster may be co-expressed, their functions appear unrelated to the PdeI-c-di-GMP pathway we are investigating. Therefore, we opted not to elaborate on these genes in our main discussion, as they do not contribute directly to our understanding of the PdeI-c-di-GMP association. However, we can include a brief mention of these genes in the manuscript, indicating their lack of relevance to the PdeI-c-di-GMP pathway. This addition will provide a more comprehensive view of the cluster's composition while maintaining our focus on the key findings related to PdeI and c-di-GMP.

We have also included the aforementioned explanations and supporting experimental data within the manuscript to clarify this important point (lines 193-217). Thank you for highlighting this apparent contradiction, allowing us to provide a more detailed explanation of our findings.

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

Overall, I found the main text of the manuscript well written and easy to understand, though too compressed in parts to fully understand the details of the work presented, some examples are outlined below. The materials and methods appeared to be less carefully compiled and could use some careful proof-reading for spelling (e.g. repeated use of "minuts" for minutes, "datas" for data) and grammar and sentence fragments (e.g. "For exponential period E. coli data." Line 333). In general, the meaning is still clear enough to be understood. I also was unable to find figure captions for the supplementary figures, making these difficult to understand.

We appreciate your careful review, which has helped us improve the clarity and quality of our manuscript. We acknowledge that some parts of the main text may have been overly compressed due to Short Report format in eLife. We have thoroughly reviewed the manuscript and expanded on key areas to provide more comprehensive explanations. We have carefully revised the Materials and Methods section to address the following: Corrected all spelling and grammatical error, including "minuts" to "minutes" and "datas" to "data". Corrected grammatical issues and sentence fragments throughout the section. We sincerely apologize for the omission of captions for the supplementary figures. We have now added detailed captions for all supplementary figures to ensure they are easily understandable. We believe these revisions address your concerns and enhance the overall readability and comprehension of our work.

General comments:

(1) To evaluate the performance of RiboD-PETRI, it would be helpful to have more details in general, particularly to do with the development of the sequencing protocol and the statistics shown. Some examples: How many reads were sequenced in each experiment? Of these, how many are mapped to the bacterial genome? How many reads were recovered per cell? Have the authors performed some kind of subsampling analysis to determine if their sequencing has saturated the detection of expressed genes? The authors show e.g. correlations between classic PETRI-seq and RiboD-PETRI for E. coli in Figure 1, but also have similar data for C. crescentus and S. aureus - do these data behave similarly? These are just a few examples, but I'm sure the authors have asked themselves many similar questions while developing this project; more details, hard numbers, and comparisons would be very much appreciated.

Thank you for your valuable feedback. To address your concerns, we have added a table in the supplementary material that clarifies the details of sequencing.

The correlation values of PETRI-seq and RiboD-PETRI data in C. crescentus are relatively good. However, the correlation values between PETRI-seq and RiboD-PETRI data in SA data are relatively less high. The reason is that the sequencing depths of RiboD-PETRI and PETRI-seq are different, resulting in much higher gene expression in the RiboD-PETRI sequencing results than in PETRI-seq, and the calculated correlation coefficient is only about 0.47. This indicates that there is some positive correlation between the two sets of data, but it is not particularly strong. This indicates that there is a certain positive correlation between these two sets of data, but it is not particularly strong. However, we have counted the expression of 2763 genes in total, and even though the calculated correlation coefficient is relatively low, it still shows that there is some consistency between the two groups of samples.

Assessment of the effect of rRNA depletion on transcriptional profiles of (A) C. crescentus (CC) and (B) S. aureus (SA) .

The Pearson correlation coefficient (r) of UMI counts per gene (log2 UMIs) between RiboD-PETRI and PETRI-seq was calculated for 4097 genes (A) and 2763 genes (B). The "ΔΔ" label represents the RiboD-PETRI protocol; The "Ctrl" label represents the classic PETRI-seq protocol we performed. Each point represents a gene.

(2) Additionally, I think it is critical that the authors provide processed read counts per cell and gene in their supplementary information to allow others to investigate the performance of their method without going back to raw FASTQ files, as this can represent a significant hurdle for reanalysis.

Thank you for your suggestion. However, it's important to clarify that reads and UMIs (Unique Molecular Identifiers) are distinct concepts in single-cell RNA sequencing. Reads can be influenced by PCR amplification during library construction, making their quantity less stable. In contrast, UMIs serve as a more reliable indicator of the number of mRNA molecules detected after PCR amplification. Throughout our study, we primarily utilized UMI counts for quantification. To address your concern about data accessibility, we have included the UMI counts per cell and gene in our supplementary materials provided above (Table S7-15. Some of the files are too large in memory and are therefore stored in GEO: GSE260458). This approach provides a more accurate representation of gene expression levels and allows for robust reanalysis without the need to process raw FASTQ files.

(3) Finally, the authors should also discuss other approaches to ribosomal depletion in bacterial scRNA-seq. One of the figures appears to contain such a comparison, but it is never mentioned in the text that I can find, and one could read this manuscript and come away believing this is the first attempt to deplete rRNA from bacterial scRNA-seq.

We have addressed this concern by including a comparison of different methods for depleting rRNA from bacterial scRNA-seq in Table S4 and make a short text comparison as follows: “Additionally, we compared our findings with other reported methods (Fig. 1B; Table S4). The original PETRI-seq protocol, which does not include an rRNA depletion step, exhibited an mRNA detection rate of approximately 5%. The MicroSPLiT-seq method, which utilizes Poly A Polymerase for mRNA enrichment, achieved a detection rate of 7%. Similarly, M3-seq and BacDrop-seq, which employ RNase H to digest rRNA post-DNA probe hybridization in cells, reported mRNA detection rates of 65% and 61%, respectively. MATQ-DASH, which utilizes Cas9-mediated targeted rRNA depletion, yielded a detection rate of 30%. Among these, RiboD-PETRI demonstrated superior performance in mRNA detection while requiring the least sequencing depth.” We have added this content in the main text (lines 110-120), specifically in relation to Figure 1B and Table S4. This addition provides context for our method and clarifies its position among existing techniques.

Detailed comments:

Line 78: the authors describe the multiplet frequency, but it is not clear to me how this was determined, for which experiments, or where in the SI I should look to see this. Often this is done by mixing cultures of two distinct bacteria, but I see no evidence of this key experiment in the manuscript.

The multiplet frequency we discuss in the manuscript is not determined through experimental mixing of distinct bacterial cultures.The PETRI-seq and mirco-SPLIT articles have also done experiments mixing the two libraries to determine the single-cell rate, and both gave good results. Our technique is derived from these two articles (mainly PETRI-seq), and the biggest difference is the difference in the later RiboD part, so we did not do this experiment separately. So the multiple frequencies here are theoretical predictions based on our sequencing results, calculated using a Poisson distribution. We have made this distinction clearer in our manuscript (lines 93-97). The method is available in Materials and Methods section (lines 520-528). The data is available in Table S2. To elaborate:

To assess the efficiency of single-cell capture in RiboD-PETRI, we calculated the multiplet frequency using a Poisson distribution based on our sequencing results

(1) Definition: In our study, multiplet frequency is defined as the probability of a non-empty barcode corresponding to more than one cell.

(2) Calculation Method: We use a Poisson distribution-based approach to calculate the predicted multiplet frequency. The process involves several steps:

We first calculate the proportion of barcodes corresponding to zero cells: . Then, we calculate the proportion corresponding to one cell: . We derive the proportion for more than zero cells: P(≥1) = 1 - P(0). And for more than one cell: P(≥2) = 1 - P(1) - P(0). Finally, the multiplet frequency is calculated as:

(3) Parameter λ: This is the ratio of the number of cells to the total number of possible barcode combinations. For instance, when detecting 10,000 cells, .

Line 94: the concept of "percentage of gene expression" is never clearly defined. Does this mean the authors detect 99.86% of genes expressed in some cells? How is "expressed" defined - is this just detecting a single UMI?

The term "percentage gene expression" refers to the proportion of genes in the bacterial strain that were detected as expressed in the sequenced cell population. Specifically, in this context, it means that 99.86% of all genes in the bacterial strain were detected as expressed in at least one cell in our sequencing results. To define "expressed" more clearly: a gene is considered expressed if at least one UMI (Unique Molecular Identifier) detected in a cell in the population. This definition allows for the detection of even low-level gene expression. To enhance clarity in the manuscript, we have rephrased the sentence as “transcriptome-wide gene coverage across the cell population”.

Line 98: The authors discuss the number of recovered UMIs throughout this paragraph, but there is no clear discussion of the number of detected expressed genes per cell. Could the authors include a discussion of this as well, as this is another important measure of sensitivity?

We appreciate your suggestion to include a discussion on the number of detected expressed genes per cell, as this is indeed another important measure of sensitivity. We would like to clarify that we have actually included statistics on the number of genes detected across all cells in the main text of our paper. This information is presented as percentages. However, we understand that you may be looking for a more detailed representation, similar to the UMI statistics we provided. To address this, we have now added a new analysis showing the number of genes detected per cell (lines 132-133, 138-139, 144-145 and 184-186, Fig. 2B, 3B and S2B). This additional result complements our existing UMI data and provides a more comprehensive view of the sensitivity of our method. We have included this new gene-per-cell statistical graph in the supplementary materials.

Figure 1B: I presume ctrl and delta delta represent the classic PETRI-seq and RiboD protocols, respectively, but this is not specified. This should be clarified in the figure caption, or the names changed.

We appreciate you bringing this to our attention. We acknowledge that the labeling in the figure could have been clearer. We have now clarified this information in the figure caption. To provide more specificity: The "ΔΔ" label represents the RiboD-PETRI protocol; The "Ctrl" label represents the classic PETRI-seq protocol we performed. We have updated the figure caption to include these details, which should help readers better understand the protocols being compared in the figure.​

Line 104: the authors claim "This performance surpassed other reported bacterial scRNA-seq methods" with a long number of references to other methods. "Performance" is not clearly defined, and it is unclear what the exact claim being made is. The authors should clarify what they're claiming, and further discuss the other methods and comparisons they have made with them in a thorough and fair fashion.

We appreciate your request for clarification, and we acknowledge that our definition of "performance" should have been more explicit. We would like to clarify that in this context, we define performance primarily in terms of the proportion of mRNA captured. Our improved method demonstrates a significantly higher rate of rRNA removal compared to other bacterial single-cell library construction methods. This results in a higher proportion of mRNA in our sequencing data, which we consider a key performance metric for single-cell RNA sequencing in bacteria. Additionally, when compared to our previous method, PETRI-seq, our improved approach not only enhances rRNA removal but also reduces library construction costs. This dual improvement in both data quality and cost-effectiveness is what we intended to convey with our performance claim.

We recognize that a more thorough and fair discussion of other methods and their comparisons would be beneficial. We have summarized the comparison in Table S4 and make a short text discussion in the main text (lines 106-120). This addition provides context for our method and clarifies its position among existing techniques.

Figure 1D: Do the authors have any explanation for the relatively lower performance of their C. crescentus depletion?

We appreciate your attention to detail and the opportunity to address this point. The lower efficiency of rRNA removal in C. crescentus compared to other species can be attributed to inherent differences between species. It's important to note that a single method for rRNA depletion may not be universally effective across all bacterial species due to variations in their genetic makeup and rRNA structures. Different bacterial species can have unique rRNA sequences, secondary structures, or associated proteins that may affect the efficiency of our depletion method. This species-specific variation highlights the challenges in developing a one-size-fits-all approach for bacterial rRNA depletion. While our method has shown high efficiency across several species, the results with C. crescentus underscore the need for continued refinement and possibly species-specific optimizations in rRNA depletion techniques. We thank you for bringing attention to this point, as it provides valuable insight into the complexities of bacterial rRNA depletion and areas for future improvement in our method.

Line 118: The authors claim RiboD-PETRI has a "consistent ability to unveil within-population heterogeneity", however the preceding paragraph shows it detects potential heterogeneity, but provides no evidence this inferred heterogeneity reflects the reality of gene expression in individual cells.

We appreciate your careful reading and the opportunity to clarify this point. We acknowledge that our wording may have been too assertive given the evidence presented. We acknowledge that the subpopulations of cells identified in other species have not undergone experimental verification. Our intention in presenting these results was to demonstrate RiboD-PETRI's capability to detect “potential” heterogeneity consistently across different bacterial species, showcasing the method's sensitivity and potential utility in exploring within-population diversity. However, we agree that without further experimental validation, we cannot definitively claim that these detected differences represent true biological heterogeneity in all cases. We have revised this section to reflect the current state of our findings more accurately, emphasizing that while RiboD-PETRI consistently detects potential heterogeneity across species, further experimental validation would be required to confirm the biological significance of the observations (lines 169-171).

Figure 1 H&I: I'm not entirely sure what I am meant to see in these figures, presumably some evidence for heterogeneity in gene expression. Are there better visualizations that could be used to communicate this?

We appreciate your suggestion for improving the visualization of gene expression heterogeneity. We have explored alternative visualization methods in the revised manuscript. Specifically, for the expression levels of marker genes shown in Figure 1H (which is Figure 2D now), we have created violin plots (Supplementary Fig. 4). These plots offer a more comprehensive view of the distribution of expression levels across different cell populations, making it easier to discern heterogeneity. However, due to the number of marker genes and the resulting volume of data, these violin plots are quite extensive and would occupy a significant amount of space. Given the space constraints of the main figure, we propose to include these violin plots as a Fig. S4 immediately following Figure 1 H&I (which is Figure 2D&E now). This arrangement will allow readers to access more detailed information about these marker genes while maintaining the concise style of the main figure.

Regarding the pathway enrichment figure (Figure 2E), we have also considered your suggestion for improvement. We attempted to use a dot plot to display the KEGG pathway enrichment of the genes. However, our analysis revealed that the genes were only enriched in a single pathway. As a result, the visual representation using a dot plot still did not produce a particularly aesthetically pleasing or informative figure.

Line 124: The authors state no significant batch effect was observed, but in the methods on line 344 they specify batch effects were removed using Harmony. It's unclear what exactly S2 is showing without a figure caption, but the authors should clarify this discrepancy.

We apologize for any confusion caused by the lack of a clear figure caption for Figure S2 (which is Figure S3D now). To address your concern, in addition to adding figure captions for supplementary figure, we would also like to provide more context about the batch effect analysis. In Supplementary Fig. S3, Panel C represents the results without using Harmony for batch effect removal, while Panel D shows the results after applying Harmony. In both panels A and B, the distribution of samples one and two do not show substantial differences. Based on this observation, we concluded that there was no significant batch effect between the two samples. However, we acknowledge that even subtle batch effects could potentially influence downstream analyses. Therefore, out of an abundance of caution and to ensure the highest quality of our results, we decided to apply Harmony to remove any potential minor batch effects. This approach aligns with best practices in single-cell analysis, where even small technical variations are often accounted for to enhance the robustness of the results.

To improve clarity, we have revised our manuscript to better explain this nuanced approach: 1. We have updated the statement to reflect that while no major batch effect was observed, we applied batch correction as a precautionary measure (lines 181-182). 2. We have added a detailed caption to Figure S3, explaining the comparison between non-corrected and batch-corrected data. 3. We have modified the methods section to clarify that Harmony was applied as a precautionary step, despite the absence of obvious batch effects (lines 492-493).

Figure 2D: I found this panel fairly uninformative, is there a better way to communicate this finding?

Thank you for your feedback regarding Figure 2D. We have explored alternative ways to present this information, using a dot plot to display the enrichment pathways, as this is often an effective method for visualizing such data. Meanwhile, we also provided a more detailed textual description of the enrichment results in the main text, highlighting the most significant findings.

Figure 2I: the figure itself and caption say GFP, but in the text and elsewhere the authors say this is a BFP fusion.

We appreciate your careful review of our manuscript and figures. We apologize for any confusion this may have caused. To clarify: Both GFP (Green Fluorescent Protein) and BFP (Blue Fluorescent Protein) were indeed used in our experiments, but for different purposes: 1. GFP was used for imaging to observe location of PdeI in bacteria and persister cell growth, which is shown in Figure 4C and 4K. 2. BFP was used for cell sorting, imaging of location in biofilm, and detecting the proportion of persister cells which shown in Figure 4D, 4F-J. To address this inconsistency and improve clarity, we will make the following corrections: 1. We have reviewed the main text to ensure that references to GFP and BFP are accurate and consistent with their respective uses in our experiments. 2. We have added a note in the figure caption for Figure 4C to explicitly state that this particular image shows GFP fluorescence for location of PdeI. 3. In the methods section, we have provided a clear explanation of how both fluorescent proteins were used in different aspects of our study (lines 326-340).

Line 156: The authors compare prices between RiboD and PETRI-seq. It would be helpful to provide a full cost breakdown, e.g. in supplementary information, as it is unclear exactly how the authors came to these numbers or where the major savings are (presumably in sequencing depth?)

We appreciate your suggestion to provide a more detailed cost breakdown, and we agree that this would enhance the transparency and reproducibility of our cost analysis. In response to your feedback, we have prepared a comprehensive cost breakdown that includes all materials and reagents used in the library preparation process. Additionally, we've factored in the sequencing depth (50G) and the unit price for sequencing (25¥/G). These calculations allow us to determine the cost per cell after sequencing. As you correctly surmised, a significant portion of the cost reduction is indeed related to sequencing depth. However, there are also savings in the library preparation steps that contribute to the overall cost-effectiveness of our method. We propose to include this detailed cost breakdown as a supplementary table (Table S6) in our paper. This table will provide a clear, itemized list of all expenses involved, including: 1. Reagents and materials for library preparation 2. Sequencing costs (depth and price per G) 3. Calculated cost per cell.

Line 291: The design and production of the depletion probes are not clearly explained. How did the authors design them? How were they synthesized? Also, it appears the authors have separate probe sets for E. coli, C. crescentus, and S. aureus - this should be clarified, possibly in the main text.

Thank you for your important questions regarding the design and production of our depletion probes. We included the detailed probe information in Supplementary Table S1, however, we didn’t clarify the information in the main text due to the constrains of the requirements of the Short Report format in eLife. We appreciate the opportunity to provide clarifications. ​

The core principle behind our probe design is that the probe sequences are reverse complementary to the r-cDNA sequences. This design allows for specific recognition of r-cDNA. The probes are then bound to magnetic beads, allowing the r-cDNA-probe-bead complexes to be separated from the rest of the library. To address your specific questions: 1. Probe Design: We designed separate probe sets for E. coli, C. crescentus, and S. aureus. Each set was specifically constructed to be reverse complementary to the r-cDNA sequences of its respective bacterial species. This species-specific approach ensures high efficiency and specificity in rRNA depletion for each organism. The hybrid DNA complex wasthen removed by Streptavidin magnetic beads. 2. Probe Synthesis: The probes were synthesized based on these design principles. 3. Species-Specific Probe Sets: You are correct in noting that we used separate probe sets for each bacterial species. We have clarified this important point in the main text to ensure readers understand the specificity of our approach. To further illustrate this process, we have created a schematic diagram showing the principle of rRNA removal and clarified the design principle in figure legend, which we have included in the figure legend of Fig. 1A.

Line 362: I didn't see a description of the construction of the PdeI-BFP strain, I assume this would be important for anyone interested in the specific work on PdeI.

Thank you for your astute observation regarding the construction of the PdeI-BFP strain. We appreciate the opportunity to provide this important information. The PdeI-BFP strain was constructed as follows: 1. We cloned the pdeI gene along with its native promoter region (250bp) into a pBAD vector. 2. The original promoter region of the pBAD vector was removed to avoid any potential interference. 3. This construction enables the expression of the PdeI-BFP fusion protein to be regulated by the native promoter of pdeI, thus maintaining its physiological control mechanisms. 4. The BFP coding sequence was fused to the pdeI gene to create the PdeI-BFP fusion construct. We have added a detailed description of the PdeI-BFP strain construction to our methods section (lines 327-334).

Reviewer #2 (Recommendations For The Authors):

(1) General remarks:

Reconsider using 'advanced' in the title. It is highly generic and misleading. Perhaps 'cost-efficient' would be a more precise substitute.

Thank you for your valuable suggestion. After careful consideration, we have decided to use "improved" in the title. Firstly, our method presents an efficient solution to a persistent challenge in bacterial single-cell RNA sequencing, specifically addressing rRNA abundance. Secondly, it facilitates precise exploration of bacterial population heterogeneity. We believe our method encompasses more than just cost-effectiveness, justifying the use of the term "advanced."

Consider expanding the introduction. The introduction does not explain the setup of the biological question or basic details such as the organism(s) for which the technique has been developed, or which species biofilms were studied.

Thank you for your valuable feedback regarding our introduction. We acknowledge our compressed writing style due to constrains of the requirements of the Short Report format in eLife. We appreciate opportunity to expand this crucial section of our manuscript, which will undoubtedly improve the clarity and impact of our manuscript's introduction.

We revised our introduction (lines 53-80) according to following principles:

(1) Initial Biological Question: We explained the initial biological question that motivated our research—understanding the heterogeneity in E. coli biofilms—to provide essential context for our technological development.

(2) Limitations of Existing Techniques: We briefly described the limitations of current single-cell sequencing techniques for bacteria, particularly regarding their application in biofilm studies.

(3) Introduction of Improved Technique: We introduced our improved technique, initially developed for E. coli.

(4) Research Evolution: We highlighted how our research has evolved, demonstrating that our technique is applicable not only to E. coli but also to Gram-positive bacteria and other Gram-negative species, showcasing the broad applicability of our method.

(5) Specific Organisms Studied: We provided examples of the specific organisms we studied, encompassing both Gram-positive and Gram-negative bacteria.

(6) Potential Implications: Finally, we outlined the potential implications of our technique for studying bacterial heterogeneity across various species and contexts, extending beyond biofilms.

(2) Writing remarks:

43-45 Reword: "Thus, we address a persistent challenge in bacterial single-cell RNA-seq regarding rRNA abundance, exemplifying the utility of this method in exploring biofilm heterogeneity.".

Thank you for highlighting this sentence and requesting a rewording. I appreciate the opportunity to improve the clarity and impact of our statement. We have reworded the sentence as: "Our method effectively tackles a long-standing issue in bacterial single-cell RNA-seq: the overwhelming abundance of rRNA. This advancement significantly enhances our ability to investigate the intricate heterogeneity within biofilms at unprecedented resolution." (lines 47-50)

49 "Biofilms, comprising approximately 80% of chronic and recurrent microbial infections in the human body..." - probably meant 'contribute to'.

Thank you for catching this imprecision in our statement. We have reworded the sentence as: "​Biofilms contribute to approximately 80% of chronic and recurrent microbial infections in the human body...​"

54-55 Please expand on "this".

Thank you for your request to expand on the use of "this" in the sentence. You're right that more clarity would be beneficial here. We have revised and expanded this section in lines 54-69.

81-84 Unclear why these species samples were either at exponential or stationary phases. The growth stage can influence the proportion of rRNA and other transcripts in the population.

Thank you for raising this important point about the growth phases of the bacterial samples used in our study. We appreciate the opportunity to clarify our experimental design. To evaluate the performance of RiboD-PETRI, we designed a comprehensive assessment of rRNA depletion efficiency under diverse physiological conditions, specifically contrasting exponential and stationary phases. This approach allows us to understand how these different growth states impact rRNA depletion efficacy. Additionally, we included a variety of bacterial species, encompassing both gram-negative and gram-positive organisms, to ensure that our findings are broadly applicable across different types of bacteria. By incorporating these variables, we aim to provide insights into the robustness and reliability of the RiboD-PETRI method in various biological contexts. We have included this rationale in our result section (lines 99-106), providing readers with a clear understanding of our experimental design choices.

86 "compared TO PETRI-seq " (typo).

We have corrected this typo in our manuscript.

94 "gene expression collectively" rephrase. Probably this means coverage of the entire gene set across all cells. Same for downstream usage of the phrase.

Thank you for pointing out this ambiguity in our phrasing. Your interpretation of our intended meaning is accurate. We have rephrased the sentence as “transcriptome-wide gene coverage across the cell population”.

97 What were the median UMIs for the 30,000 cell library {greater than or equal to}15 UMIs? Same question for the other datasets. This would reflect a more comparable statistic with previous studies than the top 3% of the cells for example, since the distributions of the single-cell UMIs typically have a long tail.

Thank you for this insightful question and for pointing out the importance of providing more comparable statistics. We agree that median values offer a more robust measure of central tendency, especially for datasets with long-tailed distributions, which are common in single-cell studies. The suggestion to include median Unique Molecular Identifier (UMI) counts would indeed provide a more comparable statistic with previous studies. We have analyzed the median UMIs for our libraries as follows and revised our manuscript according to the analysis (lines 126-130, 133-136, 139-142 and 175-180).

(1) Median UMI count in Exponential Phase E. coli:

Total: 102 UMIs per cell

Top 1,000 cells: 462 UMIs per cell

Top 5,000 cells: 259 UMIs per cell

Top 10,000 cells: 193 UMIs per cell

(2) Median UMI count in Stationary Phase S. aureus:

Total: 142 UMIs per cell

Top 1,000 cells: 378 UMIs per cell

Top 5,000 cells: 207 UMIs per cell

Top 8,000 cells: 167 UMIs per cell

(3) Median UMI count in Exponential Phase C. crescentus:

Total: 182 UMIs per cell

Top 1,000 cells: 2,190 UMIs per cell

Top 5,000 cells: 662 UMIs per cell

Top 10,000 cells: 225 UMIs per cell

(4) Median UMI count in Static E. coli Biofilm:

Total of Replicate 1: 34 UMIs per cell

Total of Replicate 2: 52 UMIs per cell

Top 1,621 cells of Replicate 1: 283 UMIs per cell

Top 3,999 cells of Replicate 2: 239 UMIs per cell

104-105 The performance metric should again be the median UMIs of the majority of the cells passing the filter (15 mRNA UMIs is reasonable). The top 3-5% are always much higher in resolution because of the heavy tail of the single-cell UMI distribution. It is unclear if the performance surpasses the other methods using the comparable metric. Recommend removing this line.

We appreciate your suggestion regarding the use of median UMIs as a more appropriate performance metric, and we agree that comparing the top 3-5% of cells can be misleading due to the heavy tail of the single-cell UMI distribution. We have removed the line in question (104-105) that compares our method's performance based on the top 3-5% of cells in the revised manuscript. Instead, we focused on presenting the median UMI counts for cells passing the filter (≥15 mRNA UMIs) as the primary performance metric. This will provide a more representative and comparable measure of our method's performance. We have also revised the surrounding text to reflect this change, ensuring that our claims about performance are based on these more robust statistics (lines 126-130, 133-136, 139-142 and 175-180).

106-108 The sequencing saturation of the libraries (in %), and downsampling analysis should be added to illustrate this point.

Thank you for your valuable suggestion. Your recommendation to add sequencing saturation and downsampling analysis is highly valuable and will help better illustrate our point. Based on your feedback, we have revised our manuscript by adding the following content:

To provide a thorough evaluation of our sequencing depth and library quality, we performed sequencing saturation analysis on our sequencing samples. The findings reveal that our sequencing saturation is 100% (Fig. 8A & B), indicating that our sequencing depth is sufficient to capture the diversity of most transcripts. To further illustrate the impact of our downstream analysis on the datasets, we have demonstrated the data distribution before and after applying our filtering criteria (Fig. S1B & C). These figures effectively visualized the influence of our filtering process on the data quality and distribution. After filtering, we can have a more refined dataset with reduced noise and outliers, which enhances the reliability of our downstream analyses.

We have also ensured that a detailed description of the sequencing saturation method is included in the manuscript to provide readers with a comprehensive understanding of our methodology. We appreciate your feedback and believe these additions significantly improve our work.

122: Please provide more details about the biofilm setup, including the media used. I did not find them in the methods.

We appreciate your attention to detail, and we agree that this information is crucial for the reproducibility of our experiments. We propose to add the following information to our methods section (lines 311-318):

"For the biofilm setup, bacterial cultures were grown overnight. The next day, we diluted the culture 1:100 in a petri dish. We added 2ml of LB medium to the dish. If the bacteria contain a plasmid, the appropriate antibiotic needs to be added to LB. The petri dish was then incubated statically in a growth chamber for 24 hours. After incubation, we performed imaging directly under the microscope. The petri dishes used were glass-bottom dishes from Biosharp (catalog number BS-20-GJM), allowing for direct microscopic imaging without the need for cover slips or slides. This setup allowed us to grow and image the biofilms in situ, providing a more accurate representation of their natural structure and composition.​"

125: "sequenced 1,563 reads" missing "with"

Thank you for correcting our grammar. We have revisd the phrase as “sequenced with 1,563 reads”.

126: "283/239 UMIs per cell" unclear. 283 and 239 UMIs per cell per replicate, respectively?

Thank you for correcting our grammar. We have revised the phrase as “283 and 239 UMIs per cell per replicate, respectively” (lines 184).

Figure 1D: Please indicate where the comparison datasets are from.

We appreciate your question regarding the source of the comparison datasets in Figure 1D. All data presented in Figure 1D are from our own sequencing experiments. We did not use data from other publications for this comparison. Specifically, we performed sequencing on E. coli cells in the exponential growth phase using three different library preparation methods: RiboD-PETRI, PETRI-seq, and RNA-seq. The data shown in Figure 1D represent a comparison of UMIs and/or reads correlations obtained from these three methods. All sequencing results have been uploaded to the Gene Expression Omnibus (GEO) database. The accession number is GSE260458. We have updated the figure legend for Figure 1D to clearly state that all datasets are from our own experiments, specifying the different methods used.

Figure 1I, 2D: Unable to interpret the color block in the data.

We apologize for any confusion regarding the interpretation of the color blocks in Figures 1I and 2D (which are Figure 2E, 3E now). The color blocks in these figures represent the p-values of the data points. The color scale ranges from red to blue. Red colors indicate smaller p-values, suggesting higher statistical significance and more reliable results. Blue colors indicate larger p-values, suggesting lower statistical significance and less reliable results. We have updated the figure legends for both Figure 2E and Figure 3E to include this explanation of the color scale. Additionally, we have added a color legend to each figure to make the interpretation more intuitive for readers.

Figure1H and 2C: Gene names should be provided where possible. The locus tags are highly annotation-dependent and hard to interpret. Also, a larger size figure should be helpful. The clusters 2 and 3 in 2C are the most important, yet because they have few cells, very hard to see in this panel.

We appreciate your suggestions for improving the clarity and interpretability of Figures 1H and 2C (which is Figure 2D, 3D now). We have replaced the locus tags with gene names where possible in both figures. We have increased the size of both figures to improve visibility and readability. We have also made Clusters 2 and 3 in Figure 3D more prominent in the revised figure. Despite their smaller cell count, we recognize their importance and have adjusted the visualization to ensure they are clearly visible. We believe these modifications will significantly enhance the clarity and informativeness of Figures 2D and 3D.​

(3) Questions to consider further expanding on, by more analyses or experiments and in the discussion:

What are the explanations for the apparently contradictory upregulation of c-di-GMP in cells expressing higher PdeI levels? How could a phosphodiesterase lead to increased c-di-GMP levels?

We appreciate the reviewer's observation regarding the seemingly contradictory relationship between increased PdeI expression and elevated c-di-GMP levels. This is indeed an intriguing finding that warrants further explanation.

PdeI was predicted to be a phosphodiesterase responsible for c-di-GMP degradation. This prediction is based on sequence analysis where PdeI contains an intact EAL domain known for degrading c-di-GMP. However, it is noteworthy that PdeI also contains a divergent GGDEF domain, which is typically associated with c-di-GMP synthesis (Fig S8). This dual-domain architecture suggests that PdeI may engage in complex regulatory roles. Previous studies have shown that the knockout of the major phosphodiesterase PdeH in E. coli leads to the accumulation of c-di-GMP. Further, a point mutation on PdeI's divergent GGDEF domain (G412S) in this PdeH knockout strain resulted in decreased c-di-GMP levels2, implying that the wild-type GGDEF domain in PdeI contributes to the maintenance or increase of c-di-GMP levels in the cell. Importantly, our single-cell experiments showed a positive correlation between PdeI expression levels and c-di-GMP levels (Response Fig. 9B). In this revision, we also constructed PdeI(G412S)-BFP mutation strain. Notably, our observations of this strain revealed that c-di-GMP levels remained constant despite increasing BFP fluorescence, which serves as a proxy for PdeI(G412S) expression levels (Fig. 4D). This experimental evidence, along with domain analysis, suggests that PdeI could contribute to c-di-GMP synthesis, rebutting the notion that it solely functions as a phosphodiesterase. HPLC LC-MS/MS analysis further confirmed that PdeI overexpression, induced by arabinose, led to an upregulation of c-di-GMP levels (Fig. 4E). These results strongly suggest that PdeI plays a significant role in upregulating c-di-GMP levels. Our further analysis revealed that PdeI contains a CHASE (cyclases/histidine kinase-associated sensory) domain. Combined with our experimental results demonstrating that PdeI is a membrane-associated protein, we hypothesize that PdeI functions as a sensor that integrates environmental signals with c-di-GMP production under complex regulatory mechanisms.

We have also included this explanation (lines 193-217) and the supporting experimental data (Fig. 4D & 4J) in our manuscript to clarify this important point. Thank you for highlighting this apparent contradiction, as it has allowed us to provide a more comprehensive explanation of our findings.

What about the rest of the genes in cluster 2 of the biofilm? They should be used to help interpret the association between PdeI and c-di-GMP.

We understand your interest in the other genes present in cluster 2 of the biofilm and their potential relationship to PdeI and c-di-GMP. After careful analysis, we have determined that the other marker genes in this cluster do not have a significant impact on biofilm formation. Furthermore, we have not found any direct relationship between these genes and c-di-GMP or PdeI. Our focus on PdeI in this cluster is due to its unique and significant role in c-di-GMP regulation and biofilm formation, as demonstrated by our experimental results. While the other genes in this cluster may be co-expressed, their functions appear to be unrelated to the PdeI and c-di-GMP pathway we are investigating. We chose not to elaborate on these genes in our main discussion as they do not contribute directly to our understanding of the PdeI and c-di-GMP association. Instead, we could include a brief mention of these genes in the manuscript, noting that they were found to be unrelated to the PdeI-c-di-GMP pathway. This would provide a more comprehensive view of the cluster composition while maintaining focus on the key findings related to PdeI and c-di-GMP.

Protein-protein interactions of marker genes in cluster 2 of 24-hour static biofilms of E coli data.

A verification is needed that the protein fusion to PdeI functional/membrane localization is not due to protein interactions with fluorescent protein fusion.

We appreciate your concern regarding the potential impact of the fluorescent protein fusion on the functionality and membrane localization of PdeI. It is crucial to verify that the observed effects are attributable to PdeI itself and not an artifact of its fusion with the fluorescent protein. To address this matter, we have incorporated a control group expressing only the fluorescent protein BFP (without the PdeI fusion) under the same promoter. This experimental design allows us to differentiate between effects caused by PdeI and those potentially arising from the fluorescent protein alone.

Our results revealed the following key observations:

(1) Cellular Localization: The GFP alone exhibited a uniform distribution in the cytoplasm of bacterial cells, whereas the PdeI-GFP fusion protein was specifically localized to the membrane (Fig. 4C).

(2) Localization in the Biofilm Matrix: BFP-positive cells were distributed throughout the entire biofilm community. In contrast, PdeI-BFP positive cells localized at the bottom of the biofilm, where cell-surface adhesion occurs (Fig 4F).

(3) c-di-GMP Levels: Cells with high levels of BFP displayed no increase in c-di-GMP levels. Conversely, cells with high levels of PdeI-BFP exhibited a significant increase in c-di-GMP levels (Fig. 4D).

(4) Persister Cell Ratio: Cells expressing high levels of BFP showed no increase in persister ratios, while cells with elevated levels of PdeI-BFP demonstrated a marked increase in persister ratios (Fig. 4J).

These findings from the control experiments have been included in our manuscript (lines 193-244, Fig. 4C, 4D, 4F, 4G and 4J), providing robust validation of our results concerning the PdeI fusion protein. They confirm that the observed effects are indeed due to PdeI and not merely artifacts of the fluorescent protein fusion.

(!) Vrabioiu, A. M. & Berg, H. C. Signaling events that occur when cells of Escherichia coli encounter a glass surface. Proceedings of the National Academy of Sciences of the United States of America 119, doi:10.1073/pnas.2116830119 (2022). https://doi.org/10.1073/pnas.2116830119

(2)bReinders, A. et al. Expression and Genetic Activation of Cyclic Di-GMP-Specific Phosphodiesterases in Escherichia coli. J Bacteriol 198, 448-462 (2016). https://doi.org:10.1128/JB.00604-15

Changes from the earlier version to the revised version

Abstract

Abstract

In contrast to mammalian cells, bacterial cells lack mRNA polyadenylated tails, presenting a hurdle in isolating mRNA amidst the prevalent rRNA during single-cell RNA-seq. This study introduces a novel method, Ribosomal RNA-derived cDNA Depletion (RiboD), seamlessly integrated into the PETRI-seq technique, yielding RiboD-PETRI. This innovative approach offers a cost-effective, equipment-free, and high-throughput solution for bacterial single-cell RNA sequencing. By efficiently eliminating rRNA reads and substantially enhancing mRNA detection rates (up to 92%), our method enables precise exploration of bacterial population heterogeneity. Applying RiboD-PETRI to investigate biofilm heterogeneity, distinctive subpopulations marked by unique genes within biofilms were successfully identified. Notably, Pdel, a marker for the cell-surface attachment subpopulation, was observed to elevate cyclic diguanylate (c-di-GMP) levels, promoting persister cell formation. Thus, we address a persistent challenge in bacterial single-cell RNA-seq regarding rRNA abundance, exemplifying the utility of this method in exploring biofilm heterogeneity. These findings advance our understanding of biofilm biology and offer insights for targeted therapeutic strategies against persistent bacterial infections.

Introduction

Earlier version: Biofilms, comprising approximately 80% of chronic and recurrent microbial infections in the human body, contain bacterial cells existing in a diverse array of physiological states 1–3. This diversity in states reflects a division of labor within the biofilm, contributing to increased resistance to various stresses4. However, the study of biofilms faces significant limitations, primarily stemming from challenges in investigating heterogeneity within a bacterial population. Single-cell RNA-seq emerges as a promising avenue for addressing this 5–12. Expending on established protocols for cell fixation and permeabilization which facilitate in-cell barcoding while avoiding cell lysis, combinatorial barcoding-based bacterial scRNA-seq techniques, such as Prokaryotic Expression profiling by Tagging RNA in situ and sequencing (PETRI-seq) 7 and Microbial Split-Pool Ligation Transcriptomics (microSPLiT)9, have been developed. Nevertheless, these methods encounter challenges in terms of low transcript recovery rates due to overwhelmingly abundant rRNA, restricting the comprehensive analysis of within-population heterogeneity. In comparison to mammalian cells, the absence of mRNA polyadenylated tails in bacteria necessitates an alternative approach for isolating mRNA (~5%) from the significantly more abundant rRNA (~95%). Here, by integrating a Ribosomal RNA derived cDNA Depletion protocol (RiboD) into a PETRI-seq, we developed RiboD-PETRI-seq that efficiently eliminates rRNA reads, thereby significantly improving mRNA detection rates and enabling exploration of within-population heterogeneity.

Revised version: Biofilms, comprising approximately 80% of chronic and recurrent microbial infections in the human body, contain bacterial cells existing in a diverse array of physiological states1–3. This diversity in states reflects a division of labor within the biofilm, contributing to increased resistance to various stresses4. However, the study of biofilms faces significant limitations, primarily stemming from challenges in investigating heterogeneity within a bacterial population. Single-cell RNA-seq emerges as a promising avenue for addressing this5–12. Expending on established protocols for cell fixation and permeabilization which facilitate in-cell barcoding while avoiding cell lysis, combinatorial barcoding-based bacterial scRNA-seq techniques, such as Prokaryotic Expression profiling by Tagging RNA in situ and sequencing (PETRI-seq) 7 and Microbial Split-Pool Ligation Transcriptomics (microSPLiT)9, have been developed. Nevertheless, these methods encounter challenges in terms of low transcript recovery rates due to overwhelmingly abundant rRNA, restricting the comprehensive analysis of within-population heterogeneity. In comparison to mammalian cells, the absence of mRNA polyadenylated tails in bacteria necessitates an alternative approach for isolating mRNA (~5%) from the significantly more abundant rRNA (~95%). Here, by integrating a Ribosomal RNA derived cDNA Depletion protocol (RiboD) into a PETRI-seq, we developed RiboD-PETRI-seq that efficiently eliminates rRNA reads, thereby significantly improving mRNA detection rates and enabling exploration of within-population heterogeneity.

Results

Earlier version: In the RiboD protocol, we designed a set of probe primers that spans all regions of the bacterial rRNA sequence. The 3’-end of the probes is specifically complementary to the rRNA-derived cDNA, while the 5’- end complements a biotin-labeled universal primer. Following template switching and RNaseH treatment on the barcoded cDNA from lysed cells to eliminate hybridized RNA, the library of probe primers and biotin-labeled universal primers is introduced to facilitate adequate hybridization. Pre-treated Streptavidin magnetic beads are then added to eliminate the hybridized rRNA-derived cDNA. The mRNA-derived cDNA remains in the supernatant and is collected for subsequent library construction and sequencing (Fig. 1A). The Multiplet frequency of this method is 1.16% – 3.35%, demonstrating the ability of RiboD-PETRI to capture transcriptomes of single cells (SI)7.

Revised version: In the RiboD protocol, we designed a set of probe primers that spans all regions of the bacterial rRNA sequence. The 3’-end of the probes is specifically complementary to the rRNA-derived cDNA, while the 5’- end complements a biotin-labeled universal primer. Following template switching and RNaseH treatment on the barcoded cDNA from lysed cells to eliminate hybridized RNA, the library of probe primers and biotin-labeled universal primers is introduced to facilitate adequate hybridization. Pre-treated Streptavidin magnetic beads are then added to eliminate the hybridized rRNA-derived cDNA. The mRNA-derived cDNA remains in the supernatant and is collected for subsequent library construction and sequencing (Fig. 1A). The Multiplet frequency of this method is 1.16% - 3.35%, demonstrating the ability of RiboD-PETRI to capture transcriptomes of single cells (Sl)7.

Removed in the revised version: Development of RiboD-PETRI and validation of its technical performance in studying population heterogeneity.

Earlier version: (A) Graphic summary of the RiboD-PETRI method illustrating the incorporation of RiboD after cell pooling and lysis in PETRI-seq. The RiboD protocol is represented by the dashed-line box. (B) Comparison of non-rRNA (tRNA, mRNA and other non-rRNA) and rRNA UMI counts ratio among different bacterial scRNA-seq methods. Data from PETRI-seq (E col/)8, MicroSPLIT-seq (E coli)9, M3-seq (E coli)11 cited from from previous studies. Error bars represent standard deviations of biological replicates. (C) Comparison of UMI counts per cell between RiboD-PETRI and PETRI at the same unsaturated sequencing depth. (D) Assessment of the effect of rRNA depletion on transcriptional profiles. The Pearson correlation coefficient (r) of UMI counts per gene (log2 UMIs) between RiboD-PETRI and PETRI was calculated for 3790 out of 4141 total genes, excluding those with zero counts in either library. Each point represents a gene. (E) Evaluation of the correlation between RiboD-PETRI data and bulk RNA-seq results. The Pearson correlation coefficient (r) of UMI counts per gene (log2 UMIs) among RiboD-PETRI data and the reads per gene (log2 reads) of bulk RNA-seq data was calculated for 3814 out of 4141 total genes, excluding those with zero counts in either library. Each point represents a gene. (F-I) RiboD-PETRI data from E coli in exponential phase (E coli 3h). (F) Distributions of mRNA UMIs captured per cell in RiboD-PETRI data of E coli in exponential phase. (G) Uniform Manifold Approximation and Projection (UMAP) visualization of E coli bacteria during the exponential phase. Data were filtered for cells with UMIs between 200 and 5,000, resulting in 1,464 cells. Each dot represents a cell. (H) Heatmap illustrating the normalized gene expression levels of marker genes in different clusters of E. coli during the exponential phase. Marker genes with relatively high expression levels are depicted in yellow, while lower expression levels are shown in purple. Each row represents a gene, and each column represents a cell. (I) Functional enrichment analysis of marker genes in cluster 2. Marker genes were selected based on screening criteria of p-value < 0.001 and log2 fold change (FC) > 0.2.

Revised version: Development of RiboD-PETRI and validation of its technical performance in studying population heterogeneity. (A) Graphic summary of the RiboD-PETRI method illustrating the incorporation of RiboD after cell pooling and lysis in PETRI-seq. The RiboD protocol is represented by the dashed-line box. (B) Comparison of non-rRNA (tRNA, mRNA and other non-rRNA) and rRNA UMI counts ratio among different bacterial scRNA-seq methods. Data from PETRI-seq (E. coli)8, MicroSPLIT-seq (E. coli)9, M3-seq (E coli)11 cited from from previous studies. Error bars represent standard deviations of biological replicates. (C) Comparison of UMI counts per cell between RiboD-PETRI and PETRI at the same unsaturated sequencing depth. (D) Assessment of the effect of rRNA depletion on transcriptional profiles. The Pearson correlation coefficient (r) of UMI counts per gene (Iog2 UMIs) between RiboD-PETRI and PETRI was calculated for 3790 out of 4141 total genes, excluding those with zero counts in either library. Each point represents a gene. (E) Evaluation of the correlation between RiboD-PETRI data and bulk RNA-seq results. The Pearson correlation coefficient (r) of UMI counts per gene (Iog2 UMIs) among RiboD-PETRI data and the reads per gene (Iog2 reads) of bulk RNA-seq data was calculated for 3814 out of 4141 total genes, excluding those with zero counts in either library. Each point represents a gene. (F-l) RiboD-PETRI data from E coli in exponential phase (E coil 3h). (F) Distributions of mRNA UMIs captured per cell in RiboD-PETRI data of E coli in exponential phase. (G) Uniform Manifold Approximation and Projection (UMAP) visualization of E coli bacteria during the exponential phase. Data were filtered for cells with UMIs between 200 and 5,000, resulting in 1,464 cells. Each dot represents a cell. (H) Heatmap illustrating the normalized gene expression levels of marker genes in different clusters of E. coli during the exponential phase. Marker genes with relatively high expression levels are depicted in yellow, while lower expression levels are shown in purple. Each row represents a gene, and each column represents a cell. (I) Functional enrichment analysis of marker genes in cluster 2. Marker genes were selected based on screening criteria of p-value < 0.001 and Iog2 fold change (FC) > 0.2.

To assess the performance of RiboD-PETRI, we conducted a comprehensive evaluation of rRNA depletion efficiency across gram-negative and gram-positive bacterial species. The results highlight a substantial enhancement in rRNA-derived cDNA depletion, with mRNA ratio increases from 8.2% to 81 % for E. coli from exponential phase, from 10% to 92% for S. aureus from stationary phase, and from 3.9% to 54% for C. crescentus from exponential phase (Fig. 1B). With equivalent sequencing depth, RiboD-PETRI demonstrates a significantly enhanced Unique Molecular Identifier (UMI) counts detection rate compared PETRI-seq alone (Fig. 1C), indicating the number of detected mRNA per cell increased prominently. Notably, this enhancement was achieved while maintaining mRNA profiles consistent with non-depleted samples (r = 0.93; Fig. 1D) and show a significant correlation with profiles from the traditional bulk RNA-seq method (r = 0.84; Fig. 1E).

Earlier version: We subsequently investigated the transcriptome coverage of RiboD-PETRI across different bacterial species. For exponential phase E. coli cells, we sequenced a library with 60,000 cells, recovering approximately 30,004 cells (50% recovery), each with ≥15 UMIs. Analysis revealed 99.86% gene expression collectively and an average of 128.8 UMIs per single cell. Top 1,000, 5,000, and 10,000 high-quality cells showed averages of 462, 259, and 193 detected UMIs, respectively (Fig. 1F). For stationary phase S. aureus cells, we sequenced a library with 30,000 cells, recovering approximately 9,982 cells (33.3% recovery), each with ≥15 UMIs. Analysis showed 99.96% gene expression collectively, and at the single-cell level, an average of 153.8 UMIs. Top 1,000, 5,000, and 8,000 high-quality cells exhibited averages of 273, 122, and 94 detected UMIs, respectively (Fig. S1A). In the exponential phase of C. crescentus cells, a library with 30,000 cells was sequenced, recovering approximately13,897 cells (46.3% recovery), each with >15 UMIs. Analysis showed 99.64% gene expression collectively, and at the singlecell level, an average of 439.7 UMIs. Top 1,000, 5,000, and 10,000 high-quality cells demonstrated averages of 2190, 662, and 225 detected UMIs, respectively (Fig. S1F). This performance surpassed other reported bacterial scRNA-seq methods7,9–11,13. Considering that this large library was sequenced with 1,508 reads per cell for E. coli, 2,372 for S. aureus, and 8,790 for C. crescentus, it is anticipated that a higher number of UMIs per cell will be detected with increased sequencing depth.

Revised version: We subsequently investigated the transcriptome coverage of RiboD-PETRI across different bacterial species. For exponential phase E. coli cells, we sequenced a library with 60,000 cells, recovering approximately 30,004 cells (50% recovery), each with ≥15 UMIs. Analysis revealed 99.86% gene expression collectively and an average of 128.8 UMIs per single cell. Top 1,000, 5,000, and 10,000 high-quality cells showed averages of 462,259, and 193 detected UMIs, respectively (Fig. 1F). For stationary phase S. aureus cells, we sequenced a library with 30,000 cells, recovering approximately 9,982 cells (33.3% recovery), each with ≥15 UMIs. Analysis showed 99.96% gene expression collectively, and at the single-cell level, an average of 153.8 UMIs. Top 1,000, 5,000, and 8,000 high-quality cells exhibited averages of 273,122, and 94 detected UMIs, respectively (Fig. S1A). In the exponential phase of C. crescentus cells, a library with 30,000 cells was sequenced, recovering approximately! 3,897 cells (46.3% recovery), each with ≥15 UMIs. Analysis showed 99.64% gene expression collectively, and at the singlecell level, an average of 439.7 UMIs. Top 1,000, 5,000, and 10,000 high-quality cells demonstrated averages of 2190,662, and 225 detected UMIs, respectively (Fig. S1F). This performance surpassed other reported bacterial scRNA-seq methods7,9–11,13. Considering that this large library was sequenced with 1,508 reads per cell for E. coli, 2,372 for S. aureus, and 8,790 for C. crescentus, it is anticipated that a higher number of UMIs per cell will be detected with increased sequencing depth.

Our results affirm RiboD-PETRI’s reliability in capturing the bacterial single-cell transcriptome, providing ample coverage and sensitivity for various species. We investigated its ability to consistently identify within-population heterogeneity across different bacterial species and growth conditions. In the exponential phase of E. coli, we recovered 1,464 cells and identified three major subpopulations (Fig. 1G), with 17 cells (1.2%) in a unique subpopulation characterized by pentose and glucuronate interconversions (Fig. 1H, I). For stationary phase S. aureus cells, we recovered 9,386 cells and found six major subpopulations (Fig. S1D), with 437 cells (4.7%) in a distinct subpopulation named cluster 4. In the stationary phase of C. crescentus cells, we recovered 5728 cells and identified four major subpopulations (Fig. S1I), with 603 cells (10.5%) in a unique subpopulation named cluster 3. These findings highlight RiboD-PETRI’s consistent ability to unveil within-population heterogeneity, crucial for understanding bacterial population complexity.

Earlier version: We next focused on exploring biological heterogeneity of a biofilm at the early stage of development by utilizing the static biofilm system14. E. coli cells were cultured in microtiter dishes overnight, adhered cells were fixed for RiboD-PETRI processing in duplicate experiments, each consisting of 1621 or 3999 cells. No significant batch effect was observed (Fig. S2C, H). Replicate 1 was sequenced 1,563 reads per cell, while replicate 2 with 2,034 reads per cell, yielding an average of 283/239 UMIs per cell (Fig. 2A). Unsupervised clustering analysis identified four major subpopulations in each replicate, with a consistently identified rare subpopulation (2.6%/2.1%) as cluster 2, driven by cell envelope genes (Fig. 2B, D). Marker genes for this cluster included yffO, lptE, rdgB, pdel, sstT, fixA, yjjG, rlml , yncD, accCand yaiA (Fig. 2C, E, F). PdeI, identified among marker genes, was predicted as a phosphodiesterase enzyme hydrolyzing c-di-GMP, a vital bacterial second messenger (Fig. 2G, H). To validate this, we created a Pdel-BFP fusion construct under the native pdeI promoter, integrated with a ratiometric c-di-GMP sensing system in E. coli. Confocal microscopy revealed PdeI as a membrane protein (Fig. 2I). Single-cell level monitoring showed cell-to-cell variability in c-di-GMP levels and PdeI expression, with a positive correlation observed (Fig. 2J). PdeI upregulated c-di-GMP synthesis rather than degradation, confirmed by both microscopy and HPLC LC-MS/MS, which showed an approximately elevenfold increase in c-di-GMP concentration in the PdeI overexpression strain compared to the control strain (Fig. 2K).

Revised version: We next focused on exploring biological heterogeneity of a biofilm at the early stage of development by utilizing the static biofilm system14. E. coli cells were cultured in microtiter dishes overnight, adhered cells were fixed for RiboD-PETRI processing in duplicate experiments, each consisting of 1621 or 3999 cells. No significant batch effect was observed (Fig. S2C, H). Replicate 1 was sequenced 1,563 reads per cell, while replicate 2 with 2,034 reads per cell, yielding an average of 283/239 UMIs per cell (Fig. 2A). Unsupervised clustering analysis identified four major subpopulations in each replicate, with a consistently identified rare subpopulation (2.6%/2.1 %) as cluster 2, driven by cell envelope genes (Fig. 2B, D). Marker genes for this cluster included yffO, IptE, rdgB, pdel, sstT, fixA, yjjG, rlml, yncD, accC and yaiA (Fig. 2C, E, F). Pdel, identified among marker genes, was predicted as a phosphodiesterase enzyme hydrolyzing c-di-GMR a vital bacterial second messenger (Fig. 2G, H). To validate this, we created a Pdel-BFP fusion construct under the native pdel promoter, integrated with a ratiometric c-di-GMP sensing system in E. coli. Confocal microscopy revealed Pdel as a membrane protein (Fig. 2I). Single-cell level monitoring showed cell-to-celI variability in c-di-GMP levels and Pdel expression, with a positive correlation observed (Fig. 2J). Pdel upregulated c-di-GMP synthesis rather than degradation, confirmed by both microscopy and HPLC LC-MS/MS, which showed an approximately elevenfold increase in c-di-GMP concentration in the Pdel overexpression strain compared to the control strain (Fig. 2K).

Removed in the revised version: RiboD-PETRI Resolves Biofilm Heterogeneity and Identifies New Marker Genes.

Earlier version: (A-D) RiboD-PETRI data from 24-hour static biofilms of E coli (E coli 24h). (A) UMI counts per cell in RiboD-PETRI data of 24-hour static biofilms of E coli. Two replicates were screened for cells with UMIs between 100 and 2000, resulting in 1621 and 3999 cells, with median UMIs per cell of 283.5 and 239, respectively. (B) UMAP visualization of 24-hour static biofilms of E. coli, revealing two small populations of heterogeneous cells in clusters 2 and 3. (C) Inferred expression levels of marker genes from 24-hour static biofilms of E. coli across different clusters. (D) Enrichment pathways for marker genes in cluster 2, selected based on screening criteria of p-value < 0.001 and log2 fold change (FC) > 0.2. (E & F) Dot plot displaying scaled expression levels of marker genes in different clusters of E. coli in exponential phase (E) and E. coli in 24- hour static biofilms (F). These genes were markers of static E. coli biofilms in cluster 2, identified with screening criteria of p-value < 0.001 and logFC > 3. Dot size represents the percentage expression of the gene in the cluster, while color indicates the average expression level normalized from 0 to 1 across all clusters for each gene. (G & H) UMAP plots showing the distribution of pdel in single-cell data of E. coli in exponential phase (G) and E. coli in 24-hour static biofilms (H). Each dot represents a cell colored by normalized expression levels of genes. (I) Subcellular localization of Pdel-GFP. (J) c-di-GMP levels (R-1 score) in E. colicells with different PdeI-GFP expression levels (low or high) in 24-hour static biofilms. c-di-GMP levels are measured using the c-di-GMP sensor system15 integrated into E. coli cells. Scale bar, 2 μm. R-1 score was determined using the fluorescent intensity of mVenusNB and mScarlet-I in the system. The fluorescent intensity is measured by sorted by fluorescence-activated cell sorting (FACS). (K) Determination of cellular concentrations of c-di-GMP by HPLC-MS/MS in cells overexpressing Pdel under the control of arabinose promoter, with 0.002% arabinose induction for 2 h (n=3). (L) Localization of Pdel-high cells in the biofilm matrix. Cells expressing Pdel-BFP under the control of the pdel native promoter were grown in a glass-bottom cell culture dish and stained with SYTO™ 24 for bacterial DNA. (M & N) Heterogeneous expression of Pdel in single-cell data of E. coli in exponential phase (M) and E. coli in 24- hour static biofilms (N). Biofilm cells with high or low expression levels of Pdel-BFP were sorted by FACS. (O) Persister counting assay using 150 pg/ml ampicillin on cells with high or low expression levels of Pdel-BFP from 24h static biofilm E. coli, sorted by FACS. (P) Time-lapse images of the persister assay observed under a microscope. Static biofilm cells of the Pdel-GFP strain were spotted on a gel pad and treated with 150 μg/ml ampicillin in LB broth. lmages were captured over 6 hours at 37°C, followed by the replacement of fresh LB broth to allow persister cell resuscitation. Scale bar, 2 μm. Error bars represent standard deviations of biological replicates. Significance was ascertained by two-tailed Student’s t test. Statistical significance is denoted as *P <0.05, **P < 0.01, ***, P<0.001, ****, P<0.0001.

Revised version: RiboD-PETRI Resolves Biofilm Heterogeneity and Identifies New Marker Genes. (A-D) RiboD-PETRI data from 24-hour static biofilms of E coli(E coli 24h). (A) UMI counts per cell in RiboD-PETRI data of 24-hour static biofilms of E coli. Two replicates were screened for cells with UM Is between 100 and 2000, resulting in 1621 and 3999 cells, with median UMIs per cell of 283.5 and 239, respectively. (B) UMAP visualization of 24-hour static biofilms of E. coli, revealing two small populations of heterogeneous cells in clusters 2 and 3. (C) Inferred expression levels of marker genes from 24-hour static biofilms of E. coli across different clusters. (D) Enrichment pathways for marker genes in cluster 2, selected based on screening criteria of p-value < 0.001 and Iog2 fold change (FC) > 0.2. (E & F) Dot plot displaying scaled expression levels of marker genes in different clusters of E. coli in exponential phase (E) and E. coli in 24- hour static biofilms (F). These genes were markers of static E. coli biofilms in cluster 2, identified with screening criteria of p-value < 0.001 and logFC > 3. Dot size represents the percentage expression of the gene in the cluster, while color indicates the average expression level normalized from 0 to 1 across all clusters for each gene. (G & H) UMAP plots showing the distribution of pdel in single-cell data of E. coli in exponential phase (G) and E. coli in 24-hour static biofilms (H). Each dot represents a cell colored by normalized expression levels of genes. (I) Subcellular localization of Pdel-GFP. (J) c-di-GMP levels (R-1 score) in E. coil cells with different Pdel-GFP expression levels (low or high) in 24-hour static biofilms. c-di-GMP levels are measured using the c-di-GMP sensor system15 integrated into E. coli cells. Scale bar, 2 μm. R-1 score was determined using the fluorescent intensity of mVenusNB and mScarlet-l in the system. The fluorescent intensity is measured by sorted by fluorescence-activated cell sorting (FACS). (K) Determination of cellular concentrations of c-di-GMP by HPLC-MS/MS in cells overexpressing Pdel under the control of arabinose promoter, with 0.002% arabinose induction for 2 h (n=3). (L) Localization of Pdel-high cells in the biofilm matrix. Cells expressing Pdel-BFP under the control of the pdel native promoter were grown in a glass-bottom cell culture dish and stained with SYTO™ 24 for bacterial DNA. (M & N) Heterogeneous expression of Pdel in single-cell data of E. coli in exponential phase (M) and E. coli in 24-hour static biofilms (N). Biofilm cells with high or low expression levels of Pdel-BFP were sorted by FACS. (0) Persister counting assay using 150 μg/ml ampicillin on cells with high or low expression levels of Pdel-BFP from 24h static biofilm E coli, sorted by FACS. (P) Time-lapse images of the persister assay observed under a microscope. Static biofilm cells of the Pdel-GFP strain were spotted on a gel pad and treated with 150 μg/ml ampicillin in LB broth. Images were captured over 6 hours at 37°C, followed by the replacement of fresh LB broth to allow persister cell resuscitation. Scale bar, 2 μm. Error bars represent standard deviations of biological replicates. Significance was ascertained by two-tailed Student’s t test. Statistical significance is denoted as *P <0.05, **P < 0.01, ***, P<0.001, ****, P<0.0001.

Earlier version: Confocal laser scanning microscopy confirmed that PdeI-positive cells, exhibiting elevated c-di-GMP levels, were predominantly located at the bottom of the static biofilm (Fig. 2L), corresponding to the region of cell-surface attachment. To investigate the association of the PdeI-high cluster with bacterial persistence in the early stage of biofilm development, PdeI-high cells were isolated by fluorescence- activated cell sorting (FACS) (Fig. 2M, N) and subjected to an ampicillin antibiotic killing assay to determine their persister frequency. Results show that the PdeI-high population produced a significant higher ratio of persister cells (~7.3%) compared to that of the whole biofilm population (~0.6%) (Fig. 2O). Time-lapse imaging during the antibiotic killing process consistently revealed that persisters mainly originated from PdeI-GFP positive cells (Fig. 2P and Movie S1). PdeI-GFP positive cells, displaying dormancy, survived ampicillin treatment for 6 hours without growth or division. Upon antibiotic removal and replacement with fresh growth medium, the PdeI-GFP positive persister cells resumed activity, elongating, dividing, and forming new microcolonies (Fig. 2P and Movie S1). These findings suggest that c-di-GMP,a molecule whose intracellular levels are upregulated by PdeI, plays a significant role in generating a persister subpopulation during the early stages of biofilm development.

Revised version: Confocal laser scanning microscopy confirmed that Pdel-positive cells, exhibiting elevated c-di-GMP levels, were predominantly located at the bottom of the static biofilm (Fig. 2L), corresponding to the region of cell-surface attachment. To investigate the association of the Pdel-high cluster with bacterial persistence in the early stage of biofilm development, Pdel-high cells were isolated by fluorescence-activated cell sorting (FACS) (Fig. 2M, N) and subjected to an ampicillin antibiotic killing assay to determine their persister frequency. Results show that the Pdel-high population produced a significant higher ratio of persister cells (~7.3%) compared to that of the whole biofilm population (~0.6%) (Fig. 2o). Time-lapse imaging during the antibiotic killing process consistently revealed that persisters mainly originated from Pdel-GFP positive cells (Fig. 2P and Movie SI). Pdel-GFP positive cells, displaying dormancy, survived ampicillin treatment for 6 hours without growth or division. Upon antibiotic removal and replacement with fresh growth medium, the Pdel-GFP positive persister cells resumed activity, elongating, dividing, and forming new microcolonies (Fig. 2P and Movie S1). These findings suggest that c-di-GMR a molecule whose intracellular levels are upregulated by Pdel, plays a significant role in generating a persister subpopulation during the early stages of biofilm development.

Discussion

Earlier version: In conclusion, we report an improved approach, called RiboD-PETRI, which offers a cost-effective ($0.0049 per cell in RiboD-PETRI, and $0.056 per cell in PETRI-seq), equipment-free, and high-throughput solution for bacterial scRNA-seq. By integrating a probe hybridization-based rRNA derived cDNA depletion protocol, our method efficiently removes rRNA reads and significantly enhances mRNA detection rates. This improvement enables us to explore and analyze within-population heterogeneity. One important advantage of RiboD is the depletion process takes place at the cDNA level, following cell pooling and lysis. Consequently, RiboD only requires customization in the library preparation step. By implementing RiboD, we effectively reduce the necessary sequencing depth and overall sequencing costs, making it an attractive option for cost-conscious researchers. Beyond serving as a proof of principle, our study exemplifies the application of RiboD-PETRI to investigate the heterogeneity within biofilms, specifically delving into the early stages of biofilm development. This showcases the versatility and efficacy of RiboD-PETRI in exploring complex biological systems. The ability to uncover hidden variations within bacterial populations, as demonstrated in our biofilm analysis, underlines the potential impact of RiboDPETRI on advancing our understanding of microbial behavior and population dynamics.

Revised version: In conclusion, we report an improved approach, called RiboD-PETRI, which offers a cost-effective ($0.0049 per cell in RiboD-PETRI, and $0,056 per cell in PETRI-seq), equipment-free, and high-throughput solution for bacterial scRNA-seq. By integrating a probe hybridization-based rRNA derived cDNA depletion protocol, our method efficiently removes rRNA reads and significantly enhances mRNA detection rates. This improvement enables us to explore and analyze within-population heterogeneity. One important advantage of RiboD is the depletion process takes place at the cDNA level, following cell pooling and lysis. Consequently, RiboD only requires customization in the library preparation step. By implementing RiboD, we effectively reduce the necessary sequencing depth and overall sequencing costs, making it an attractive option for cost-conscious researchers. Beyond serving as a proof of principle, our study exemplifies the application of RiboD-PETRI to investigate the heterogeneity within biofilms, specifically delving into the early stages of biofilm development. This showcases the versatility and efficacy of RiboD-PETRI in exploring complex biological systems. The ability to uncover hidden variations within bacterial populations, as demonstrated in our biofilm analysis, underlines the potential impact of RiboD-PETRI on advancing our understanding of microbial behavior and population dynamics.

Materials and methods

Resource availability

Further information and requests for resources and reagents should be directed to and will be fulfilled by the lead contact, Yingying Pu (yingyingpu@whu.edu.cn).

Materials availability

Plasmids generated in this study are available from the lead contact upon request.

Data and code availability

Sequencing data, the processing data and original code have been deposited to GEO repository (Uploading).

Bacterial strains and growth conditions

Earlier version: The bacterial strains used in this study included Escherichia coli strains MG1655, Caulobacter crescentus NA1000 and Staphylococcus aureus strain ATCC 25923 were included. Cultures of E. coli were grown in Luria Broth (LB) medium. Caulobacter crescentus strain NA1000 was grown in peptone yeast extract (PYE) medium. And Staphylococcus aureus strain ATCC 25923 was grown in Mueller-Hinton Broth (MHB) medium. All bacterial strains were routinely grown at 37°C and 220 rpm. To maintain plasmids, when necessary, media were supplemented with chloramphenicol (25 μg/ml). For arabinose-induction system expression experiments, 0.002% arabinose was supplemented in the medium.

Revised version: The bacterial strains used in this study included Escherichia coli strains MG1655, Caulobacter crescentus NAI 000 and Staphylococcus aureus strain ATCC 25923 were included. Cultures of E. coli were grown in Luria Broth (LB) medium. Caulobacter crescentus strain NAI 000 was grown in peptone yeast extract (PYE) medium. And Staphylococcus aureus strain ATCC 25923 was grown in Mueller-Hinton Broth (MHB) medium. All bacterial strains were routinely grown at 37°C and 220 rpm. To maintain plasmids, when necessary, media were supplemented with chloramphenicol (25 μg/ml). For arabinose-induction system expression experiments, 0.002% arabinose was supplemented in the medium.

Strains construction.

Earlier version: The construction of recombinant plasmids was performed using the 2 × MultiF Seamless Assembly Mix (ABclonal, RK21020). To label the Pdel gene, pdel, along with its native promoter, was fused with either gfp or bfp and cloned into the pBAD backbone. For gene overexpression, the target genes of interest were cloned into either a p15A ori plasmid or a pUC ori plasmid, under the control of the Arabinose-induction system. For the c-di-GMP sensor (Snapgene: #182291), the plasmid origin was replaced with the p15A ori.

Revised version: The construction of recombinant plasmids was performed using the 2 × MultiF Seamless Assembly Mi× (ABclonal, RK21020). To label the Pdel gene, pdel, along with its native promoter, was fused with either gfp or bfp and cloned into the pBAD backbone. For gene overexpression, the target genes of interest were cloned into either a p15A ori plasmid or a pUC ori plasmid, under the control of the Arabinose-induction system. For the c-di-GMP sensor (Snapgene: #182291), the plasmid origin was replaced with the p15A ori.

RiboD-PETRI

Cell preparation.

Earlier version: E. coli MG1655 cells were cultured overnight and subsequently diluted at a ratio of 1:100 into fresh LB medium and grown statically for 24h at 37°C.For 3h exponential period E. coli sample, E. coliMG1655 cells were grown overnight and then diluted 1:100 into fresh LB medium and grown for 3 h at 37°C and 220 rpm. Caulobacter crescentus strain NA1000 cells were grown overnight and then diluted 1:100 into fresh Mueller-Hinton Broth (MHB) medium and grown for 9h at 37°C and 220 rpm. And Staphylococcus aureus strain ATCC 25923 cells were grown overnight and then diluted 1:100 into fresh peptone yeast extract (PYE) medium and grown for 3 h at 37°C and 220 rpm. All the culture was vigorously shaken using a vortex, and the cells were then centrifuged at 5,000g for 2 minutes at 4°C. The pellet was resuspended in 2 ml of ice-cold 4% formaldehyde (F8775, Millipore Sigma, diluted into PBS). This suspension was rotated at 4°C for 16 hours.

Revised version: E. coli MG1655 cells were cultured overnight and subsequently diluted at a ratio of 1:100 into fresh LB medium and grown statically for 24h at 37°C. For 3h exponential period E. coli sample, E. coli MG1655 cells were grown overnight and then diluted 1:100 into fresh LB medium and grown for 3 h at 37°C and 220 rpm. Caulobacter crescentus strain NAI 000 cells were grown overnight and then diluted 1:100 into fresh Mueller-Hinton Broth (MHB) medium and grown for 9h at 37°C and 220 rpm. And Staphylococcus aureus strain ATCC 25923 cells were grown overnight and then diluted 1:100 into fresh peptone yeast extract (PYE) medium and grown for 3 h at 37°C and 220 rpm. All the culture was vigorously shaken using a vortex, and the cells were then centrifuged at 5,000g for 2 minutes at 4°C. The pellet was resuspended in 2 ml of ice-cold 4% formaldehyde (F8775, Millipore Sigma, diluted into PBS). This suspension was rotated at 4°C for 16 hours.

Cell permeabilization.

Earlier version: 1 ml of fixed cells were centrifuged at 5,000g for 5 minuts at 4°C, followed by resuspension in 1 ml washing buffer (100 mM Tris-HCl pH7.0, 0.02 U/μl SUPERase-In RNase Inhibitor, AM2696, Invitrogen). After another centrifugation at 5,000g for 5 minuts at 4°C, the supernatant was removed. The pellet was then resuspended in 250 μl permeabilization buffer (0.04% Tween-20 in PBS-RI, PBS with 0.01 U/μl SUPERase-In RNase Inhibitor) and incubated on ice for 3 minuts. 1 ml cold PBS-RI was added, and the cells were centrifuged at 5,000g for 5 minuts at 4°C, followed by resuspension in 250 μl Lysozyme Mix (Dissolving 250 μg/ml Lysozyme or 5 μg/ml Lysostaphin for S. aureus in TEL-RI buffer, comprising 100 mM Tris, pH 8.0 (AM9856, Invitrogen), 50 mM EDTA (AM9261, Invitrogen), and 0.1 U/μl SUPERase In RNase Inhibitor). The samples were incubated at 37°C and mixed gently every minute. Then cold PBS-RI (1ml) was added immediately, and cells were centrifuged at 5,000g for 5 min at 4°C. The cells underwent another wash with 1 ml cold PBS-RI. Subsequently, cells were resuspended in 40 μl DNaseI-RI buffer (4.4 μl 10×reaction buffer, 0.2 μl SUPERase In RNase inhibitor, 35.4 μl H2O), followed by addition of 4 μl DNaseI (AMPD1, Millipore Sigma). The samples were incubated for 30 min at room temperature and mixed gently every 5 minutes. 4 μl Stop Solution was added, and the samples were incubated for 10 minutes at 50°C with gentle mixing every minuts. Following centrifugation at 5,000g for 10 minutes at 4°C, cells were washed twice with 0.5 ml cold PBS-RI. Finally, cells were resuspended in 200 μl cold PBS-RI, and their count and integrity were assessed using the ACEA NovoCyte flow cytometer with a 100×oil immersion lens.

Revised version: 1 ml of fixed cells were centrifuged at 5,000g for 5 minuts at 4°C, followed by resuspension in 1 ml washing buffer (100 mM Tris-HCI pH7.0,0.02 U/μl SUPERase-ln RNase Inhibitor, AM2696, Invitrogen). After another centrifugation at 5,000g for 5 minuts at 4°C, the supernatant was removed. The pellet was then resuspended in 250 μl permeabilization buffer (0.04% Tween-20 in PBS-RI, PBS with 0.01 U/μl SUPERase-ln RNase Inhibitor) and incubated on ice for 3 minuts. 1 ml cold PBS-RI was added, and the cells were centrifuged at 5,000g for 5 minuts at 4°C, followed by resuspension in 250 μl Lysozyme Mix (Dissolving 250 μg/ml Lysozyme or 5 μg/ml Lysostaphin for S. aureus in TEL-RI buffer, comprising 100 mM Tris, pH 8.0 (AM9856, Invitrogen), 50 mM EDTA (AM9261, Invitrogen), and 0.1 U/μl SUPERase In RNase Inhibitor). The samples were incubated at 37°C and mixed gently every minute. Then cold PBS-RI (1 ml) was added immediately, and cells were centrifuged at 5,000g for 5 min at 4°C. The cells underwent another wash with 1 ml cold PBS-RI. Subsequently, cells were resuspended in 40 μl DNasel-RI buffer (4.4 μl 10xreaction buffer, 0.2 μl SUPERase In RNase inhibitor, 35.4 μl H2O), followed by addition of 4 μl DNasel (AMPD1, Millipore Sigma). The samples were incubated for 30 min at room temperature and mixed gently every 5 minutes. 4 μl Stop Solution was added, and the samples were incubated for 10 minutes at 50°C with gentle mixing every minuts. Following centrifugation at 5,000g for 10 minutes at 4°C, cells were washed twice with 0.5 ml cold PBS-RI. Finally, cells were resuspended in 200 μl cold PBS-RI, and their count and integrity were assessed using the ACEA NovoCyte flow cytometer with a 100× oil immersion lens.

Primer preparation.

Earlier version: For round 1 RT, round 2 and round 3 ligation reactions, all primers design and preparation as previously described 7. All primers were purchased from Sangon Biotech (Table S1). For ligation primers preparation, mixtures were prepared as follows: 31.1 μl each R2 primer (100 μM), 28.5 μl SB83 (100 μM) and 21.4 μl H2O, were splitted to 2.24 μl for one sample. Mixtures containing 63.2 μl each R3 primer (70 μM) and 58 μl SB8 (70 μM), were splitted to 3.49 μl for one sample. Before use, ligation primers were incubated as follows: 95 °C for 3 min, then decreasing the temperature to 20 °C at a ramp speed of −0.1 °C s-1, 37 °C for 30 min. For blocking mix preparation, 50 μl primer SB84 (400 μM) and 80 μl primer SB81 (400 μM) were incubated as follows: 94°C for 3 min, then decreasing the temperature to 25 °C at a ramp speed of −0.1 °C s-1, 4°C for keeping. Round 2 blocking primers were mixed as follows: 37.5 μl 400 μM SB84, 37.5 μl 400 μM SB85, 25 μl 10x T4 ligase buffer, 150 μl H2O. Round 3 blocking primers were mixed as follows: 72 μl 400 μM SB81,72 μl 400 μM SB82, 120 μl 10×T4 ligase buffer, 336 μl H2O, 600 μL 0.5 M EDTA.

Revised version: For round 1 RT, round 2 and round 3 ligation reactions, all primers design and preparation as previously described 7. All primers were purchased from Sangon Biotech (Table S1). For ligation primers preparation, mixtures were prepared as follows: 31.1 μl each R2 primer (100 μM), 28.5 μl SB83 (100 μM) and 21.4 μl H2O, were splitted to 2.24 μl for one sample. Mixtures containing 63.2 μl each R3 primer (70 μM) and 58 μl SB8 (70 μM), were splitted to 3.49 μl for one sample. Before use, ligation primers were incubated as follows: 95 °C for 3 min, then decreasing the temperature to 20 °C at a ramp speed of -0.1 °C s-1,37 °C for 30 min. For blocking mix preparation, 50 μl primer SB84 (400 μM) and 80 μl primer SB81 (400 μM) were incubated as follows: 94°C for 3 min, then decreasing the temperature to 25 °C at a ramp speed of -0.1 °C s-1,4°C for keeping. Round 2 blocking primers were mixed as follows: 37.5 μl 400 μM SB84,37.5 μl 400 μM SB85,25 μl 10× T4 ligase buffer, 150 μl H2O. Round 3 blocking primers were mixed as follows: 72 μl 400 μM SB81,72 μl 400 μM SB82,120 μl 10× T4 ligase buffer, 336 μl H2O, 600 μl 0.5 M EDTA.

Round 1 RT reaction.

Earlier version: About 3×107 cells were introduced into an RT reaction mix composed of 240 μl 5×RT buffer, 24 μl dNTPs (N0447L, NEB), 12 μl SUPERase In RNase Inhibitor and 24 μl Maxima H Minus Reverse Transcriptase (EP0753, Thermo Fisher Scientific). Nuclease-free water was added to achieve a total reaction volume of 960 μl, and the mixture was thoroughly mixed by vortexing. Subsequently, 8 μl of the reaction mixture was dispensed into each well of a 96-well plate, where 2 μl of each RT primer had been added previously. The sealed 96-well plate was inverted repeatedly for thorough mixing, followed by a brief spin. The plate was then incubated as follows: 50 °C for 10 min, 8 °C for 12 s, 15 °C for 45 s, 20 °C for 45 s, 30 °C for 30 s, 42 °C for 6 min, 50 °C for 16 min, and finally held at 4 °C. After the RT process, all 96 reactions were pooled into one tube. 75 μl of 0.5% Tween-20 was added, and the reactions were incubated on ice for 3 minutes. Cells were centrifuged at 7,000g for 10 minutes at 4°C and then resuspended in 0.4 ml PBS-RI. Thirty-two microliters of 0.5% Tween-20 was added, and the cells underwent centrifugation at 7,000g for 10 minutes at 4°C.

Revised version: About 3×107 cells were introduced into an RT reaction mix composed of 240 μl 5× RT buffer, 24 μl dNTPs (N0447L, NEB), 12 μl SUPERase In RNase Inhibitor and 24 μl Maxima H Minus Reverse Transcriptase (EP0753, Thermo Fisher Scientific). Nuclease-free water was added to achieve a total reaction volume of 960 μl, and the mixture was thoroughly mixed by vortexing. Subsequently, 8 μl of the reaction mixture was dispensed into each well of a 96-well plate, where 2 μl of each RT primer had been added previously. The sealed 96-well plate was inverted repeatedly for thorough mixing, followed by a brief spin. The plate was then incubated as follows: 50 °C for 10 min, 8 °C for 12 s, 15 °C for 45 s, 20 °C for 45 s, 30 °C for 30 s, 42 °C for 6 min, 50 °C for 16 min, and finally held at 4 °C. After the RT process, all 96 reactions were pooled into one tube. 75 μl of 0.5% Tween-20 was added, and the reactions were incubated on ice for 3 minutes. Cells were centrifuged at 7,000g for 10 minutes at 4°C and then resuspended in 0.4 ml PBS-RL Thirty-two microliters of 0.5% Tween-20 was added, and the cells underwent centrifugation at 7,000g for 10 minutes at 4°C.

Round 2 ligation reaction.

Cells were resuspended in 500 μl 1× T4 ligase buffer, followed by the addition of 107.5 μl PEG8000,37.5 μl 10× T4 ligase buffer, 16.7 μl SUPERase In RNase Inhibitor, 5.6 μl BSA, and 27.9 μl T4 ligase (M0202L, NEB). The reaction solution was thoroughly mixed by vortexing. Subsequently, 5.76 μl of the reaction mixture was dispensed into each well of a 96-well plate, where 2.24 μl of each round 2 ligation primer had been added previously. The sealed 96-well plate was inverted repeatedly for thorough mixing and then subjected to a short spin. The plate was incubated at 37°C for 45 minutes. Following this, 2 μl of round 2 blocking mix was added to each well and incubated at 37°C for an additional 45 minutes. All 96 reactions were pooled into one tube after incubation.

Round 3 ligation reaction.

Earlier version: A mixture comprising 89 μl H2O, 26 μl PEG8000, 46 μl 10×T4 ligase buffer, and 12.65 μl T4 ligase was prepared and thoroughly mixed by vortexing. Subsequently, 8.51 μl of the reaction mixture was dispensed into each well of a 96-well plate, where 3.49 μl of each round 3 ligation primer had been added previously. The sealed 96-well plate was inverted repeatedly for thorough mixing and then subjected to a brief spin. The plate was incubated at 37°C for 45 minutes. Following this, 10 μl of round 3 blocking mix was added to each well and incubated at 37°C for an additional 45 minutes. All 96 reactions were combined into one tube after incubation.

Revised version: A mixture comprising 89 μl H20,26 μl PEG8000,46 μl 10× T4 ligase buffer, and 12.65 μl T4 ligase was prepared and thoroughly mixed by vortexing. Subsequently, 8.51 μl of the reaction mixture was dispensed into each well of a 96-well plate, where 3.49 μl of each round 3 ligation primer had been added previously. The sealed 96-well plate was inverted repeatedly for thorough mixing and then subjected to a brief spin. The plate was incubated at 37°C for 45 minutes. Following this, 10 μl of round 3 blocking mix was added to each well and incubated at 37°C for an additional 45 minutes. All 96 reactions were combined into one tube after incubation.

Cells lysis.

Earlier version: 42 μl of 0.5% Tween-20 was added, and cells were centrifuged at 7,000g for 10 minutes at 4°C. The cells underwent two washes using 200 μl TEL-RI containing 0.01% Tween-20, each time centrifuged at 7,000g for 10 minutes at 4°C. Subsequently, cells were resuspended in 30 μl TEL-RI buffer. Cell counting and integrity checks were performed using the ACEA NovoCyte flow cytometer with a 100× oil immersion lens. A moderate amount of cells was then added to the lysis buffer (50 mM Tris pH 8.0, 25 mM EDTA, 200 mM NaCl, 0.5% Triton X-100), and 5 μl of proteinase K (AM2548, Invitrogen) was introduced. Samples were incubated at 55°C for 60 minutes and gently mixed every minute.

Revised version: 42 μl of 0.5% Tween-20 was added, and cells were centrifuged at 7,000g for 10 minutes at 4°C. The cells underwent two washes using 200 μl TEL-RI containing 0.01 % Tween-20, each time centrifuged at 7,000g for 10 minutes at 4°C. Subsequently, cells were resuspended in 30 μl TEL-RI buffer. Cell counting and integrity checks were performed using the ACEA NovoCyte flow cytometer with a 100× oil immersion lens. A moderate amount of cells was then added to the lysis buffer (50 mM Tris pH 8.0,25 mM EDTA, 200 mM NaCI, 0.5% Triton X-100), and 5 μl of proteinase K (AM2548, Invitrogen) was introduced. Samples were incubated at 55°C for 60 minutes and gently mixed every minute.

Library construction.

Earlier version: To facilitate template switching, lysates were purified with VAHTS DNA Clean Beads (N411, Vazyme) at a ratio of 2.0×, and cDNA was eluted in 12 μl of water. The purified cDNA was then combined with 4 μl of 5× RT buffer, 1 μl of dNTPs (N0447L, NEB), 0.5 μl of SUPERase In RNase Inhibitor, 0.5 μl of Maxima H Minus Reverse Transcriptase, and 0.5 μl of the TSO primer (100 mM, Table S1). This reaction solution underwent incubation as follows: 25°C for 30 min, 42°C for 90 min, 85°C for 5 min, and then held at 4°C. Subsequently, 1 μl of RNaseH was added, and the reaction solution was incubated at 37°C for 30 min. The cDNA was purified once again with VAHTS DNA Clean Beads at a ratio of 2.0× and eluted in 13 μl of H2O. The integrity of the cDNA was assessed using primers TSO-2 and R1 or R2 or R3 by qPCR (Table S1).

Revised version: To facilitate template switching, lysates were purified with VAHTS DNA Clean Beads (N411, Vazyme) at a ratio of 2.0×, and cDNA was eluted in 12 μl of water. The purified cDNA was then combined with 4 μl of 5× RT buffer, 1 μl of dNTPs (N0447L, NEB), 0.5 μl of SUPERase In RNase Inhibitor, 0.5 μl of Maxima H Minus Reverse Transcriptase, and 0.5 μl of the TSO primer (100 mM, Table S1). This reaction solution underwent incubation as follows: 25°C for 30 min, 42°C for 90 min, 85°C for 5 min, and then held at 4°C. Subsequently, 1 μl of RNaseH was added, and the reaction solution was incubated at 37°C for 30 min. The cDNA was purified once again with VAHTS DNA Clean Beads at a ratio of 2.0× and eluted in 13 μl of H2O. The integrity of the cDNA was assessed using primers TSO-2 and RI or R2 or R3 by qPCR (Table S1).

Earlier version: Ribosomal RNA derived cDNA depletion (RiboD)

Revised version: Ribosomal RNAderived cDNA depletion (RiboD).

Earlier version: We developed a set of cDNA probe primers to selectively deplete r-cDNA (Table S1). These probe primers possess the ability to specifically hybridize with r-cDNA and also hybridize with a biotin-labeled universal primer. In the reaction, 5 μl of r-cDNA probe primers (10 μM), 2.5 μl of 10× hybridization buffer (Tris-HCl pH 8.0 100 mM, NaCl 500 mM, EDTA pH 8.0 10 mM), and 5 μl of biotin primer (10 μM) were added to 12.5 μl of purified cDNA. The reaction solution underwent incubation as follows: 95°C for 2 min, followed by a temperature decrease to 20°C at a ramp speed of −0.1°C s-1, and then held at 37°C for 30 min. Subsequently, 20 μl of Streptavidin magnetic beads (BEAVER, 22307) was washed twice using 1 ml of 1× B&W buffer (Tris-HCl pH 7.5 10 mM, EDTA 1 mM, NaCl 1M, Tween-20 0.05%) and resuspended in 25 μl of 2× B&W buffer. Twenty-five microliters of washed Streptavidin magnetic beads were added to 25 μl of annealed cDNA. The reaction solution was incubated at room temperature for 30 min with gentle mixing per minute. Following this, the reaction solution tube was placed into a magnetic stand to collect the supernatant. The cDNA depleted of r-cDNA was purified using VAHTS DNA Clean Beads at a ratio of 2.0× and eluted in 12.5 μl of H2O. The depletion of r-cDNA could be repeated using the above protocol, and ultimately, the cDNA was eluted in 20 μl of H2O.

Revised version: We developed a set of cDNA probe primers to selectively deplete r-cDNA (Table S1). These probe primers possess the ability to specifically hybridize with r-cDNA and also hybridize with a biotin-labeled universal primer. In the reaction, 5 μl of r-cDNA probe primers (10 μM), 2.5 μl of 10× hybridization buffer (Tris-HCI pH 8.0 100 mM, NaCI 500 mM, EDTA pH 8.0 10 mM), and 5 μl of biotin primer (10 μM) were added to 12.5 μl of purified cDNA. The reaction solution underwent incubation as follows: 95°C for 2 min, followed by a temperature decrease to 20°C at a ramp speed of -0.1 °C s-1, and then held at 37°C for 30 min. Subsequently, 20 μl of Streptavidin magnetic beads (BEAVER, 22307) was washed twice using 1 ml of 1 × B&W buffer (Tris-HCI pH 7.5 10 mM, EDTA 1 mM, NaC11M, Tween-20 0.05%) and resuspended in 25 μl of 2× B&W buffer. Twenty-five microliters of washed Streptavidin magnetic beads were added to 25 μl of annealed cDNA. The reaction solution was incubated at room temperature for 30 min with gentle mixing per minute. Following this, the reaction solution tube was placed into a magnetic stand to collect the supernatant. The cDNA depleted of r-cDNA was purified using VAHTS DNA Clean Beads at a ratio of 2.0× and eluted in 12.5 μl of H2O. The depletion of r-cDNA could be repeated using the above protocol, and ultimately, the cDNA was eluted in 20 μl of H2O.

Library amplification and sequencing.

Earlier version: To the 20 μl cDNA solution, the following components were added: 2.4 μl R3 primer (10 mM, Table S1), 2.4 μl TSO-2 primer (10 mM, Table S1), 40 μl 2× KAPA HIFI mix (KAPA, 2602), 1.6 μl Sybr green (25×), 0.8 μl MgCl2 (0.1 M), and 12.8 μl H2O. This PCR reaction solution was placed in a thermocycler and incubated with the following parameters: 98 °C for 45 s, followed by cycling of 98 °C for 15 s, 60 °C for 30 s, and 72 °C for 60 s. Cycling continued on a qPCR machine until the reaction approached saturation. PCR products were then purified using VAHTS DNA Clean Beads at a ratio of 0.9× and eluted in 25 μl of H2O. Finally, the purified PCR products underwent end repair and adaptor ligation using the VAHTS Universal DNA Library Prep Kit for Illumina V3 (Vazyme, ND607).

Revised version: To the 20 μl cDNA solution, the following components were added: 2.4 μl R3 primer (10 mM, Table SI), 2.4 μl TSO-2 primer (10 mM, Table SI), 40 μl 2× KAPA HIFI mix (KAPA, 2602), 1.6 μl Sybr green (25x), 0.8 μl MgCI2 (0.1 M), and 12.8 μl H2O. This PCR reaction solution was placed in a thermocycler and incubated with the following parameters: 98 °C for 45 s, followed by cycling of 98 °C for 15 s, 60 °C for 30 s, and 72 °C for 60 s. Cycling continued on a qPCR machine until the reaction approached saturation. PCR products were then purified using VAHTS DNA Clean Beads at a ratio of 0.9× and eluted in 25 μl of H20. Finally, the purified PCR products underwent end repair and adaptor ligation using the VAHTS Universal DNA Library Prep Kit for Illumina V3 (Vazyme, ND607).

Bulk RNA-seq library construction

Total RNA of the samples was extracted utilizing the Bacteria RNA Extraction Kit (R403-01, Vazyme). Subsequently, the RNA underwent mRNA enrichment (N407, Vazyme), fragmentation, cDNA synthesis, and library preparation using the VAHTSTM Total RNA-seq (H/M/R) Library Prep Kit for Illumina® (NR603, Vazyme).

Bioinformatics analysis Methods

Single-Cell Analysis.

Earlier version: The sequencing data underwent processing into matrices using scripts and a pipeline as previously described7 in Python 2.7.15, with some modifications (the detailed original code and all the data were deposited in the GEO repository). After the count tables were made, subsequent analysis of single-cell data was conducted using Seurat package (version 4.3.0; http://satijalab.org/seurat/) in R (https://www.r-project.org/). Since there are two replicates of E coli static biofilm, these two datas need to be merged into one SeuratObject and removed batch effects. However, the samples for exponential period E. coli, S. aureus and C. crescentus only had one sample, so they did not need this process. At the beginning of doing the scRNA-seq analysis, we screened the datas of all samples. For preprocessing of E coli static biofilm data, cells were filtered with UMI per cell more than 100 and less than 2000 for replicate 1 and replicate 2 to obtain 1621 and 3999 cells, respectively. For exponential period E coli data. The data was screened for cells with UMIs greater than 200 and less than 5000 to obtain 1464 cells. The screening criteria of S. aureus were cells with UMIs greater than 15 and less than 1000 and genes greater than 30 (1000>UMIs>15, gene counts>30). The screening criteria of C. crescentus were cells with UMIs greater than 200 and less than 5000 and gene counts greater than 30 (5000>UMIs>200, gene counts>30). After screening, all the datas were normalized using a scale factor of 10000 through a global-scaling normalization method called “LogNormalize”. Highly variable features were then identified, returning 500 features per dataset. Then we combine the data of the two replicates of E coli static biofilm into a single SeuratObject by FindIntegrationAnchors and IntegrateData functions. Then all the datas underwent scaling using the ScaleData function, followed by dimension reduction through principal component analysis. After principal component analysis, we removed batch effect by RunHarmony of the two replicates of E coli static biofilm datas. Then a graph-based clustering approach was employed in all datas to identify clusters of gene expression programs using the Louvain algorithm (Seurat 4.3.0). The dims we choosed were 6. And the resolution was 0.3 for C. crescentus and S. aureus or 0.4 for E coli datas. Marker genes for each cluster were computed using the Wilcoxon Ranksum test. Specifically, marker genes for each cluster were initially obtained using the FindMarkers function of Seurat. Then we performed pathway enrichment analysis of marker gene by clusterProfiler function16 within R.

Revised version: The sequencing data underwent processing into matrices using scripts and a pipeline as previously described7 in Python 2.7.15, with some modifications (the detailed original code and all the data were deposited in the GEO repository). After the count tables were made, subsequent analysis of single-cell data was conducted using Seurat package (version 4.3.0; http://satijalab.org/seurat/) in R (https://www.r-project.org/). Since there are two replicates of E. coli static biofilm, these two datas need to be merged into one SeuratObject and removed batch effects. However, the samples for exponential period E. coli, S. aureus and C. crescentus only had one sample, so they did not need this process. At the beginning of doing the scRNA-seq analysis, we screened the datas of all samples. For preprocessing of E. coli static biofilm data, cells were filtered with UMI per cell more than 100 and less than 2000 for replicate 1 and replicate 2 to obtain 1621 and 3999 cells, respectively. For exponential period E. coli data. The data was screened for cells with UMIs greater than 200 and less than 5000 to obtain 1464 cells. The screening criteria of S. aureus were cells with UMIs greater than 15 and less than 1000 and genes greater than 30 (1000>UMIs>15, gene counts>30). The screening criteria of C. crescentus were cells with UMIs greater than 200 and less than 5000 and gene counts greater than 30 (5000>UMIs>200, gene counts>30). After screening, all the datas were normalized using a scale factor of 10000 through a global-scaling normalization method called “LogNomnalize”. Highly variable features were then identified, returning 500 features per dataset. Then we combine the data of the two replicates of E. coli static biofilm into a single SeuratObject by FindlntegrationAnchors and IntegrateData functions. Then all the datas underwent scaling using the ScaleData function, followed by dimension reduction through principal component analysis. After principal component analysis, we removed batch effect by RunHarmony of the two replicates of E. coli static biofilm datas. Then a graph-based clustering approach was employed in all datas to identify clusters of gene expression programs using the Louvain algorithm (Seurat 4.3.0). The dims we choosed were 6. And the resolution was 0.3 for C. crescentusand S. aureus or 0.4 for E coli datas. Marker genes for each cluster were computed using the Wilcoxon Ranksum test. Specifically, marker genes for each cluster were initially obtained using the FindMarkers function of Seurat. Then we performed pathway enrichment analysis of marker gene by clusterProfiler function16 within R.

Comparison of scRNA-seq with Bulk RNA-Seq.

Earlier version: The bulk RNA-seq clean data reads were mapped to the E coli MG1655 k12 genome (EnsemblBacteria Taxonomy ID: 511145) using the BWA aligner software (v0.7.17-r1188, https://github.com/lh3/bwa.git). Converting sam files to bam files using samtool (v1.9). The mapping results were counted by featureCounts (v2.0.1, https://github.com/topics/featurecounts) to generate expression results. Single-cell and bulk transcriptomes of E coli were compared by computing the Pearson correlation of log2 reads per gene of bulk RNA-seq and log2 UMI per gene of scRNA-seq.

Revised version: The bulk RNA-seq clean data reads were mapped to the E coli MG1655 k12 genome (EnsemblBacteria Taxonomy ID: 511145) using the BWA aligner software (vO.7.17-r1188, https://github.com/lh3/bwa.git). Converting sam files to bam files using samtool (v1.9). The mapping results were counted by featurecounts (v2.0.1, https://github.com/topics/featurecounts) to generate expression results. Single-cell and bulk transcriptomes of E. coli were compared by computing the Pearson correlation of Iog2 reads per gene of bulk RNA-seq and Iog2 UMI per gene of scRNA-seq.

Flow cytometry and FACS analysis

Earlier version: All samples were measured on a Beckman CytoFLEX SRT flow cytometer with a 70 μm nozzle in which normal saline was used as sheath fluid. PdeI-BFP or c-di-GMP sensor labeled strain in the 24 h static biofilm growth phase was washed and resuspended in sterile PBS. Microorganisms were identified by FSC (forward scatter) and SSC (side scatter) parameters. Cells were sorted into different groups based on their fluorescence intensity (PB450 for BFP, FITC for mVenusNB, ECD for mScarlet-I). The results were analyzed by FlowJo V10 software (Treestar, Inc.).

Revised version: All samples were measured on a Beckman CytoFLEX SRT flow cytometer with a 70 μM nozzle in which normal saline was used as sheath fluid. Pdel-BFP or c-di-GMP sensor labeled strain in the 24 h static biofilm growth phase was washed and resuspended in sterile PBS. Microorganisms were identified by FSC (forward scatter) and SSC (side scatter) parameters. Cells were sorted into different groups based on their fluorescence intensity (PB450 for BFR FITC for mVenusNB, ECD for mScarlet-l). The results were analyzed by FlowJo VI0 software (Treestar, Inc.).

Antibiotic killing and persister counting assay

Cells sorted from FACS were suspended in fresh LB broth containing 150 μg/ml ampicillin and incubated for an additional 3 hours at 37°C with continuous shaking at 220 rpm. Following this, the cells were plated on LB plates for colony-forming unit (CFU) counts before the ampicillin challenge, and the plates were incubated overnight at 37°C. After the ampicillin challenge, cells were washed with PBS buffer and plated again for CFU counts. The persister ratio was defined as the ratio of CFU after ampicillin challenge to CFU before ampicillin challenge. Averages and standard deviations presented are representative of three biological replicates.

Microscopy

Bright-field and fluorescence imaging.

Earlier version: Inverted microscopes (Nikon Eclipse Ti2 and Leica Stellaris 5 WLL) were utilized, with illumination provided by different lasers: 405 nm for BFP, 488 nm for GFP, respectively. The fluorescence emission signal was captured by an sCMOS camera (pco.edge 4.2 bi). Filter sets tailored to each fluorophores spectrum were employed. Image analysis was conducted using ImageJ software (Fiji). For c-di-GMP sensor analysis, the ratio of mVenusNB to mScarlet-I (R) exhibited a negative correlation with the concentration of c-di-GMP. Therefore, R-1 displayed a positive correlation with the concentration of c-di-GMP.

Revised version: Inverted microscopes (Nikon Eclipse Ti2 and Leica Stellaris 5 WLL) were utilized, with illumination provided by different lasers: 405 nm for BFR 488 nm for GFR respectively. The fluorescence emission signal was captured by an sCMOS camera (pco.edge 4.2 bi). Filter sets tailored to each fluorophore’s spectrum were employed. Image analysis was conducted using Imaged software (Fiji). For c-di-GMP sensor analysis, the ratio of mVenusNB to mScarlet-l (R) exhibited a negative correlation with the concentration of c-di-GMP. Therefore, R-1 displayed a positive correlation with the concentration of c-di-GMP.

Time-lapse imaging.

To observe antibiotic killing and bacteria resuscitation processes, cells labeled with Pdel-GFP in the 24-hour static growth phase were collected, washed twice with PBS buffer, and imaged on a gel pad containing 3% low melting temperature agarose in PBS [REF]. The gel pad was prepared in the center of the FCS3 chamber as a gel island. The cells were then observed under bright field or epifluorescence illumination. Subsequently, the gel pad was surrounded by LB broth containing 150 μg/ml ampicillin for 6 hours at 35°C. Fresh LB broth was flushed in, and the growth medium was refreshed every 3 hours, allowing cells to recover sufficiently.

Determination of c-di-GMP concentration by HPLC-MS/MS

Earlier version: Determination of c-di-GMP concentration by HPLC-MS/MS involved growing MG1655 Δara pBAD::pdeI and MG1655 Δara to mid-exponential growth phase, followed by induction with 0.002% arabinose. After 2 hours of incubation, cells were harvested and washed with PBS. The washed cells were rapidly frozen using liquid nitrogen. Simultaneously, another portion of washed cells was stained with SYTOTM24 and quantified using flow cytometry. c-di-GMP concentration detection was conducted by Wuhan Lixinheng Technology Co. Ltd. through high-pressure liquid chromatography-tandem mass spectrometry (HPLC-MS/MS). All strains were assayed in biological triplicates, and measured values were converted into intracellular c-di-GMP concentrations (pg) per cell.

Revised version: Determination of c-di-GMP concentration by HPLC-MS/MS involved growing MG1655 Δara pBAD::pdel and MG1655 Δara to mid-exponential growth phase, followed by induction with 0.002% arabinose. After 2 hours of incubation, cells were harvested and washed with PBS. The washed cells were rapidly frozen using liquid nitrogen. Simultaneously, another portion of washed cells was stained with SYT0TM24 and quantified using flow cytometry. c-di-GMP concentration detection was conducted by Wuhan Lixinheng Technology Co. Ltd. through high-pressure liquid chromatography-tandem mass spectrometry (HPLC-MS/MS). All strains were assayed in biological triplicates, and measured values were converted into intracellular c-di-GMP concentrations (pg) per cell.

Quantification and statistical analysis

Statistical analysis was conducted using GraphPad Prism 9 software for Windows. Significance was determined by a two-tailed Student’s t-test. Error bars indicate the standard deviations of the mean from a minimum of three independent experiments. A significance threshold of P < 0.05 was applied. Asterisks were used to denote significant differences (*, P < 0.05; **, P < 0.01; ***, P < 0.001; ****, P < 0.0001).